Conformational dynamics of a nicotinic receptor neurotransmitter binding site

Author Response

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1 (Public Review):

Weaknesses:

The weaknesses are the brevity of the simulations, the concomitant lack of scope of the simulations, the lack of depth in the analysis, and the incomplete relation to other relevant work.

A 1 µs simulation of CCh (Video 1, part 2) shows that m3 (ACHA) is stable, throughout. The DG comparisons, in silico versus in vitro, indicate that 200 ns simulations are sufficient to identify LA versus HA conformational populations. Figure 6-table supplement 1 shows distances. New citations have been added.

Reviewer #2 (Public Review):

Weaknesses:

After carrying out all-atom molecular dynamics, the authors revert to a model of binding using continuum Poisson-Boltzmann, surface area, and vibrational entropy. The motivations for and limitations associated with this approximate model for the thermodynamics of binding, rather than using modern atomistic MD free energy methods (that would fully incorporate configurational sampling of the protein, ligand, and solvent) could be provided. Despite this, the authors report a correlation between their free energy estimates and those inferred from the experiment. This did, however, reveal shortcomings for two of the agonists. The authors mention their trouble getting correlation to experiment for Ebt and Ebx and refer to up to 130% errors in free energy. But this is far worse than a simple proportional error, because -24 Vs -10 kcal/mol is a massive overestimation of free energy, as would be evident if the authors were to instead express results in terms of KD values (which would have an error exceeding a billion fold). The MD analysis could be improved with better measures of convergence, as well as a more careful discussion of free energy maps as a function of identified principal components, as described below. Overall, however, the study has provided useful observations and interpretations of agonist binding that will help understand pentameric ligand-gated ion channel activation.

The objective of the calculations was to identify structural populations, not to estimate binding free energies. We knew the actual LA and HA energies (for all 4 agonists) from real-world electrophysiology experiments. We conclude that the simple PBSA method worked as a tool for identification because the calculated efficiencies match those from experiments (Figure 4B, Figure 4-Source Data 1). We discuss the mismatches in absolute G in the Results and Discussion. Methods for estimating experimental binding free energies are described in a cited, eLife companion paper. The G ratio relates to agonist efficiency.

Main points:

Regarding the choice of model, some further justification of the reduced 2 subunit ECD-only model could be given. On page 5 the authors argue that, because binding free energies are independent of energy changes outside the binding pocket, they could remove the TMD and study only an ECD subunit dimer. While the assumption of distant interactions being small seems somewhat reasonable, provided conformational changes are limited and localised, how do we know the packing of TMD onto the ECD does not alter the ability of the alpha-delta interface to rearrange during weak or strong binding? They further write that "fluctuations observed at the base of the ECD were anticipated because the TMD that offers stability here was absent.". As the TMD-ECD interface is the "gating interface" that is reshaped by agonist binding, surely the TMD-ECD interface structure must affect binding. It seems a little dangerous to completely separate the agonist binding and gating infrastructure, based on some assumption of independence. Given the model was only the alpha and delta subunits and not the pentamer with TMD, I am surprised such a model was stable without some heavy restraints. The authors state that "as a further control we carried out MD simulation of a pentamer docked with ACh and found similar structural changes at the binding pocket compared to the dimer." Is this sufficient proof of the accuracy of the simplified model? How similar was the model itself with and without agonist in terms of overall RMSD and RMSD for the subunit interface and the agonist binding site, as well as the free energy of binding to each model to compare?

The statement that distant interactions are small is not an "assumption", but rather a conclusion based on data. Mutant cycle analysis of 83 pairs shows (with a few exceptions) non-additivity of free energy change prevails only with separations <~15 A (Fig.3 in Gupta et al 2017). Regardless, the adequacy of dimers and convergence by 200 ns are supported by the calculated and experimental agonist efficiencies match (Figure 4B) and the 1 ms simulation (Video 1 part 2). Apo 200ns simulation of the ECD dimer is now added (Figure 2-figure supplement 2) and the dimer interface seems to be adequate (stable).

Although the authors repeatedly state that they have good convergence with their MD, I believe the analysis could be improved to convince us. On page 8 the authors write that the RMSD of the system converged in under 200 ns of MD. However, I note that the graph is of the entire ECD dimer, not a measure for the local binding site region. An additional RMSD of local binding site would be much more telling. You could have a structural isomerisation in the site and not even notice it in the existing graph. On page 9 the authors write that the RMSF in Figure S2 showed instability mainly in loops C and F around the pocket. Given this flexibility at the alpha-delta interface, this is why collecting those regions into one group for the calculation of RMSD convergence analysis would have been useful. They then state "the final MD configuration (with CCh) was well-aligned with the CCh-bound cryo-EM desensitized structure (7QL6)... further demonstrating that the simulation had converged." That may suggest a change occurred that is in common with the global minimum seen in cryo EM, which is good, but does not prove the MD has "converged". I would also rename Figure S3 accordingly.

The description is now changed to “aligns well” with desensitized structure (7QL6.PDB)”. RMSD of not just the binding pocket but the whole ECD dimer is well aligned with first apo (m1) and with desensitized state (m3).

The authors draw conclusions about the dominant states and pathways from their PCA component free energy projections that need clarification. It is important first to show data to demonstrate that the two PCA components chosen were dominant and accounted for most of the variance. Then when mapping free energy as a function of those two PCA components, to prove that those maps have sufficient convergence to be able to interpret them. Moreover, if the free energies themselves cannot be used to measure state stability (as seems to be the case), that the limitations are carefully explained. First, was PCA done on all MD trajectories combined to find a common PC1 & PC2, or were they done separately on each simulation? If so, how similar are they? The authors write "the first two principal components (PC-1 and PC-2) that capture the most pronounced C. displacements". How much of the total variance did these two components capture? The authors write the changes mostly concern loop C and loop F, but which data proves this? e.g. A plot of PC1 and PC2 over residue number might help.

The PCA analyses have been enriched. Figure 3-Source Data 1. shows the dominance of PC1 and PC2. Because the binding energy match was sufficient to identify affinity states, we did not explore additional PCs. Residue-wise PC1 and PC2 analysis and comparison with RMSF are in Figure 2-figure supplement 2. PC1 and PC2 both correlate with fluctuations in loops C and F. Overlap analysis in different runs is shown in Figure 3-figure supplement 1. Lower variance in a particular region of the PCA landscape indicates that the system frequently visits these states, suggesting stability (a preference for these conformations).

The authors map the -kTln rho as a free energy for each simulation as a function of PC1 & PC2. It is important to reveal how well that PC1-2 space was sampled, and how those maps converged over time. The shapes of the maps and the relative depths of the wells look very different for each agonist. If the maps were sampled well and converged, the free energies themselves would tell us the stabilities of each state. Instead, the authors do not even mention this and instead talk about "variance" being the indicator of stability, stating that m3 is most stable in all cases. While I can believe 200ns could not converge a PC1-2 map and that meaningful delta G values might not be obtained from them, the issue of lack of sampling must be dealt with. On page 12 they write "Although the bottom of the well for 3 energy minima from PCA represent the most stable overall conformation of the protein, they do not convey direct information regarding agonist stability or orientation". The reasons why not must be explained; as they should do just that if the two order parameters PC1 and PC2 captured the slowest degrees of freedom for binding and sampling was sufficient. The authors write that "For all agonists and trajectories, m3 had the least variance (was most stable), again supporting convergence by 200 ns." Again the issue of actual free energy values in the maps needs to be dealt with. The probabilities expressed as -kTln rho in kcal/mol might suggest that m2 is the most stable. Instead, the authors base stability only on variance (I guess breadth of the well?), where m3 may be more localised in the chosen PC space, despite apparently having less preference during the MD (not the lowest free energy in the maps).

The motivations and justifications for the use of approximate PBSA energetics instead of atomistic MD free energies should be dealt with in the manuscript, with limitations more clearly discussed. Rather than using modern all-atom MD free energy methods for relative or absolute binding free energies, the author selects clusters from their identified states and does Poisson-Boltzmann estimates (electrostatic, vdW, surface area, vibrational entropy). I do believe the following sentence does not begin to deal with the limitations of that method: "there are limitations with regard to MM-PBSA accurately predicting absolute binding free energies (Genheden & Ryde, 2015; Hou et al., 2011) that depends on the parameterization of the ligand (Oostenbrink et al., 2004)." What are the assumptions and limitations in taking continuum electrostatics (presumably with parameters for dielectric constants and their assignments to regions after discarding solvent), surface area (with its assumptions and limitations), and of course assuming vibration of a normal mode can capture entropy. On page 30, regarding their vibrational entropy estimate, they write that the "entropy term provides insights into the disorder within the system, as well as how this disorder changes during the binding process". It is important that the extent of disorder captured by the vibrational estimate be discussed, as it is not obvious that it has captured entropy involving multiple minima on the system's true 3N-dimensional energy surface, and especially the contribution from solvent disorder in bound Vs dissociated states.

As discussed above, errors in the free energy estimates need to be more faithfully represented, as fractional errors are not meaningful. On page 21 the authors write "The match improved when free energy ratios rather than absolute values were compared." But a ratio of free energies is not a typical or expected measure of error in delta G. They also write "For ACh and CCh, there is good agreement between.Gm1 and GLA and between.Gm3 and GHA. For these agonists, in silico values overestimated experimental ones only by ~8% and ~25%. The agreement was not as good for the other 2 agonists, as calculated values overestimated experimental ones by ~45%(Ebt) and ~130% (Ebt). However, the fractional overestimation was approximately the same for GLA and GHA." See the above comment on how this may misrepresent the error. On page 21 they write, in relation to their large fractional errors, that they "do not know the origin of this factor but speculate that it could be caused by errors in ligand parameterization". However the estimates from the PBSA approach are, by design, only approximate. Both errors in parameterisation (and their likely origin) and the approximate model used, need discussion.

Again, the goal of calculating binding free energy was to identify structural correspondence to LA and HA and not to obtain absolute binding free energy values. Along with the least variance (distribution) for the principle component for m3, it also had the highest binding free energy. An association of m1 to LA and m3 to HA was done after comparing them to experimental values (efficiencies). This comparison not only validates our approach but also underscores the utility of PBSA in supplementing MD and PCA analyses with broader energetics perspectives.

Reviewer #3 (Public Review):

Weaknesses:

Although the match in simulated vs experimental energies for two ligands was very good, the calculated energies for two other ligands were significantly different than the experiment. It is unclear to what extent the choice of method for the energy calculations influenced the results. See above.

A control simulation, such as for an apo site, is lacking. Figure 2-figure supplement 2. shows the results of 200 ns MD simulations of the apo structure (n=2).

Reviewer #4 (Public Review):

Weaknesses:

Timescales (200 ns) do not capture global rearrangements of the extracellular domain, let alone gating transitions of the channel pore, though this work may provide a launching point for more extended simulations. A more general concern is the reproducibility of the simulations, and how representative states are defined. It is not clear whether replicates were included in principal component analysis or subsequent binding energy calculations, nor how simulation intervals were associated with specific states.

We are interested eventually in using MD to study the full isomerization, but these investigations are for the future and likely will involve full length pentamers and longer timescales. However, in response to this query we have in the Discussion raised this issue and offer speculations. See above, PCA has be compared between replicates (Figure 3-figure supplement 1).

Structural analysis largely focuses on snapshots, with limited direct evidence of consistency across replicates or clusters. Figure legends and tables could be clarified.

Snapshots and distance measurements (Figure 6-table supplement 1) were extracted from m1, m2 and m3 plateau regions of trajectories. Incorporated in the legend.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

This study gives interesting insights into the possible dynamics of ligand binding in ACh receptors and establishes some prerequisites for necessary and urgent further work. The broad interest in this receptor class means this work will have some reach.

Suggestions:

(1) I found the citation of relevant literature to be rather limited. In the following paper, the agonist glutamate was shown to bind in two different orientations, and also to convert. These are much longer simulations than what is presented here (nearly 50 µs), which allowed a richer view of conformational changes and ligand binding dynamics in the AMPA Receptor. Albert Lau has published similar work on NMDA, delta, and kainate receptors, including some of it in eLife. Perhaps the authors could draw some helpful comparisons with this work.

Yu A et al. (2018) Neurotransmitter Funneling Optimizes Glutamate Receptor Kinetics. Neuron

Likewise, the comparison to a similar piece of work on glycine receptors (not cited, https://pubs.acs.org/doi/10.1021/bi500815f) could be instructive. Several similar computational techniques were used, and interactions observed (in the simulations) between the agonist and the receptor were tested in the context of wet experiments. In the absence of an equivalent process in this paper (no findings were tested using an orthogonal approach, only compared against known results, from perhaps a narrow spectrum of papers), we have to view the major findings of the paper (docking in cis that leads to a ligand somersault) with some hesitancy.

The Gharpure 2019 paper is cited in the context of the delta subunit but this paper was about a3b4 neuronal nicotinic receptors. This could be tidied up. Also, the simulations from that paper could be used as an index of the stability of the HA state (if ligand orientation is being cited as transferrable, other observations could be too).

New citations have been added. It is difficult to generalize from Yu A and Yu R eta al, because in neither study was the ligand orientation associated with LA versus HA binding energy.

(2) "To start, we associated the agonist orientation in the hold end states as cis in AC-LA versus trans in AC-HA."

I think this a valid start, but one is left with the feeling that this is all we have and the validity of the starting state is not tested. What was really shown here? Is the docking reliable? What evidence can the authors summon for the ligand orientation that they use as a starting structure? In addition to docking energies, the match between PBSA and electrophysiology Gs and temporal sequence (m1-m2-m3) support the assignment.

Given that these simulations cover a circumscribed part of the binding process, I think the limitations should be acknowledged. Indeed the authors do mention a number of remaining open questions.

Paragraphs regarding 'catch' have been added to the Discussion.

(3) Results around line 90. Hypothetical structures and states that were determined from Markov analyses are discussed as if they are well understood and identified. Plausible though these are, I think the text should underline at least the source of such information. In these simulations, a further intermediate has been identified.

The model in Figure 1B was first published in 2012 and has been used and extended over the intervening years. In our lab, catch-and-hold is standard. We have published many papers (in top journals), plus reviews, regarding this scheme. We made presentations that are on Youtube. Here, at the end of the Introduction we now cite a new review article (Biophysical Journal, 2024). I am not sure what more we can do to raise awareness regarding catch and hold.

(4) The figures are dense and could be better organised. Figure 2 is key but has a muddled organization. The placement of the panel label (C) makes it look like the top row (0 ns) is part of (A). Panel B- what is shown in the oval inset (not labeled or in legend). Why not show more than one view, perhaps a sequence of time points? It is confusing to change the colour of the loops in (C). Please show the individual values in D.

Figure 2 has been redone.

(5) A lot is made of the aK145 salt bridge with aD200 and the distances - but I didn't see any measurements, or time course. This part is vague to the point of having no meaning ("bridge tightening").

We present a Table of distance measurements in the SI (Figure 6-table supplement 1).

Reviewer #2 (Recommendations For The Authors):

All main comments have been given in the above review. There are a few other minor comments below.

The 4 agonists examined were acetylcholine (ACh), carbamylcholine (CCh), epibatidine (Ebt), and epiboxidine (Ebx). Could the choices be motivated for the reader?

New in Methods: the agonists are about the same size yet represent different efficiency classes (citation to companion eLife paper). One of our (unmet) objectives was to understand the structural correlates of agonist efficiency.

The authors write that state structures generated in the MD simulation were identified by aligning free energy values with those from experiments. It would be good to explain to the reader, in the introduction, how LA and HA free energies were extracted from experiments, rather than relying on them to read older papers.

In the Introduction, we say that to get G, just measure an equilibrium constant and take the log. We think it is excessive to explain in detail in this paper how to measure the equilibrium binding constants (several methods suffice). However, we have added in Methods our basic approach: measure KLA and L2 by using electrophysiology, and compute KHA from the thermodynamic cycle using L0. We think this paper is best understood in the context of its companion, also in eLife.

In all equilibrium equations of the type A to B (e.g. on page 5), rather than using "=" signs it would be much better to use equilibrium reversible arrow symbols.

It is incorporated.

Reviewer #3 (Recommendations For The Authors):

(1) Although the match in simulated vs experimental energies for two ligands was very good, the calculated energies for Ebt and Ebx were significantly different than the experiment. Are there any alternative methods for calculating binding energies from the MD simulations that could be readily compared to?

See above. We did not use more sophisticated energy calculations because we already knew the answers. Our objective was to identify states, not to calculate energies.

(2) It would be nice to see control simulations of an apo site to ensure that the conformational changes during the MD are due to the ligands and not an artifact of the way the system is set up. I am primarily asking about this as the simulation of the isolated ECDs for the binding site interface seems like it may be unhappy without the neighboring domains that would normally surround it. On that note, was the protein constrained in any way during the MD?

Apo simulation results are presented in Figure 2-figure supplement 2. The dimer interface seems to be adequate (stable).

(3) Figure 4A-B: Should the colors for m1 and m3 be reversed?

Colors have been changed and a bar chart has been added.

Reviewer #4 (Recommendations For The Authors):

(1) Although simulations are commendably run in triplicate, it is difficult in some places to discern their consistency.

(1a) Table S1 provides important quantification of deviations in different replicates and with different agonists. Please confirm that the reported values are accurate. All values reported for the epibatidine system are identical to those reported for carbamylcholine, which seems statistically improbable. Similarly, runs 1 and 3 with epiboxidine seem identical to one another, and runs 1 and 2 with acetylcholine are nearly the same.

Figure 2-Source Data 1 has been corrected.

(1b) In reference to Figure S3, the authors comment that the simulated system (one replicate with carbamylcholine) converges within 0.5 Å RMSD of a desensitized experimental structure. This seems amazing; please specify over what atoms this deviation was calculated and with reference to what alignment. It would be interesting to know the reproducibility of this remarkable convergence in additional replicates or with other ligands; for example, Figure 5 indicates that loop C transitions to a lesser extent in the context of epibatidine than other agonists.

The comparison was for the entire dimer ECD; 0.5 Å is the result. It may be worthwhile to pursue this remarkable convergence, but not in this paper. Here, we are concerned with identifying ACLA and ACHA. Similarity between ACHA and AD structures is for a different study.

(1c) For principal-component and subsequent analyses, it appears that only one trajectory was considered for each system. Please clarify whether this is the case; if so, a rationale for the selection would be helpful, and some indication of how reproducible other replicates are expected to be.

We have added new PCA results (Results, Figure 3-figure supplement 1) that show comparable principal components in other replicates.

(2) Figure 3 shows free energy landscapes defined by principal components of fluctuation in Cα positions.

(2a) Do experimental structures (e.g. PDB IDs 6UWZ, 7QL6u) project onto any of these landscapes in informative ways?

6UWZ.pdb matches well with the apo (7QKO.pdb), comparable to m1, and 7QL6.pdb with the m3.

(2b) Please indicate the meaning of colored regions in the righthand panels.

The color panels in the top left panel indicate the colored regions in the righthand panel also, which is indicative of direction and magnitude of changes with PC1 and PC2.

(2c) Please also check the legend; do the porcupine plots really "indicate the direction and magnitude of changes between PC1 and PC2," or rather between negative and positive values of each principal component?

It indicates the direction and magnitude of changes with PC1 and PC2.

(3) It would be helpful to clarify how trajectory segments were assigned to specific minima, particularly m2 and m3.

(3a) Please verify the timeframes associated with the m2 minima, reported as "20-50 ns [with acetylcholine], 50-60 ns [with carbamylcholine], 60-100 ns [with epibatidine, and] 100-120 ns [with epiboxidine]." It seems improbable that these intervals would interleave so precisely in independent systems. Furthermore, the intervals associated with acetylcholine and epiboxidine do not appear to correspond to the m2 regions indicated in Figure S8.

Times are given in Figure 4-Source Data 1 and Figure 3-figure supplement 2. The m2 classification is based on loop displacement as well as agonist orientation. For all agonists, the selection was strictly from PCA and cluster analysis.

(3b) The text (and legend to Figure 3) indicate that 180+ ns of each trajectory was assigned to m3, which seems surprisingly consistent. However, Figure S5 indicates this minimum is more variable, appearing at 160 ns with acetylcholine but at 186 ns with carbamylcholine. Please clarify.

see above: the selection was from PCA and cluster analysis. Times are in Figure 3-figure supplement 2 and also in Figure 4-Source Data 1 (none in Fig. 3 legend).

(3c) Figures 5, 6, S6, and S7 illustrate structural features of free-energy minima in each ligand system. Please clarify what is shown, e.g. a representative snapshot, centroid, or average structure from a particular prominent cluster associated with a given minimum.

They are all representative snapshots (now in Methods). Snapshots and distance measurements (Figure 6-table supplement 1) were extracted from m1, m2 and m3 plateau regions of trajectories.

(4) Figure S4 helpfully shows the behavior of a pentameric control system; however, some elements are unclear.

(4a) The 2.5-6.5 Å jump in RMSD at ~40 ns seems abrupt; can it be clarified whether this corresponds to a transition to either m2 or m3 poses, or to another feature of e.g. alignment?

Figure 2-figure supplement 4 left bottom is just the ligand. The jump is the flip, m1 to m2.

(4b) It seems difficult to reconcile the apparently bimodal distribution of states with the proposed 3-state model. Into which RMSD peak would the m2 intermediate fall?

The simulations are only to 100 ns, where we found a complete flip of the agonist represented in the histograms. This confirmed that dimer showed similar pattern as the pentamer. In depth analysis was only done only on dimers.

(4c) The top panel is labeled "Com" with a graphical legend indicating "ACh." Does this indicate the ligand or, as described in the text legend, "the pentamer" (i.e. the receptor)? For both panels, please verify whether they are calculated on the basis of center-of-mass, heavy atoms, Cα, etc.

"Com" (for complex) has been changed to system (protein+ligand).

(5) Minor concerns:

(5a) In Figures 1 and S3, correct the PDB references (6UWX and 7QL7 are not nAChRs).

They are now corrected.

(5b) In Figure 4, do all panels represent mean {plus minus} standard deviation calculated across all cluster-frames reported in Table 1?

Yes.

Also check the graphical legend in panel A: presumably the red bars correspond to m1/LA, and the blue to m3/HA?

Corrected

(5c) In the legend to Figure S1, please clarify that panel B is reproduced from Indurthi & Auerbach 2023.

This figure has been deleted.

(5d) As indicated in Figure S2, it seems surprising that the RMSF is so apparently low at the periphery, where the subunits should contact neighbors in the extracellular domain; how might the authors account for this? Specify whether these results apply to all replicates of each system.

The redness in the periphery for all four systems indicates the magnitude of fluctuation. As we focus on the orthosteric site, we highlight the loops around the agonist binding pocket and kept other regions 75% transparent. We now include Apo simulations and the dimer appears to be stable even without an agonist present.

(5e) Within each minimum in Figure S5, three "prominent" clusters appear to be colored (by heteroatom) with carbons in cyan, pink, and yellow respectively. If this is correct, note these colors in the text legend.

Colors have been added to the legend.

(5f) In Figure S6, note in the legend that key receptor sidechains are shown as spheres, with the ligand as balls-and-sticks, and that ligand conformations in both low- and high-affinity complexes are shown in both receptor states for comparison.

This is now added in the legend.

(5g) The legend to Figure S6 also notes "The agonists are as in Fig S4," but that figure contains a single replicate of a different system; please check this reference.

This has been updated to Figure 5.

(5h) In Figure S8, the colors in the epibatidine system appear different from the others.

The colors are the same for m1, m2 and m3 in all systems including epibatidine.

(5i) In Table 1, does "n clusters" indicate the number of simulation frames included in the three prominent clusters chosen for MM-PBSA analysis? Perhaps "n frames" would be more clear.

It was a good suggestion. It has now been changed to ‘n frames’

(5j) Pg 24-ln 453 presumably should read "...that separate it from m1 and m3..."

This sentence is now changed in the discussion.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Agonists activate receptors by interacting more strongly with active versus resting conformations of their target sites. For each ligand, the strong-weak binding free energy difference sets efficacy and the weak/strong ratio sets efficiency. We performed molecular dynamics simulations to explore the conformational dynamics of a nicotinic receptor neurotransmitter binding site in the weak→strong structural transition. The alternative conformations were identified by comparing calculated and experimental binding free energies for 4 agonists. In weak→strong, the agonist rotates about its cationic center (a ‘flip’), loop C moves in (a ‘flop’) to reposition αY190 to form a water-mediated cross-subunit hydrogen bond with the ligand. The flop restructures the aromatic core, and the flip increases van der Waals interactions to generate a more compact, hydrophobic and stable pocket. The simulations reveal a transient intermediate state as well as changes in a salt bridge that may distinguish agonists.

Revised version: Agonists turn on receptors because they provide net favorable binding energy to active versus resting conformations of their target sites. We used simulations to explore conformational dynamics of the weak→strong binding transition at the Torpedo α–δ nicotinic acetylcholine receptor orthosteric site. Using 4 agonists, the alternative site conformations were identified in trajectories generated from a single starting structure by matching binding energies calculated in silico with those measured experimentally in vitro. The weak→strong transition starts with a rotation of the agonist about its cationic center (‘flip’), followed by a downward displacement of loop C that repositions αY190 (‘flop’), followed by formation of H-bonds between the ligand, a structural water and the δ subunit loop E backbone (‘fix’). The result is a compact, hydrophobic and stable pocket with higher affinity for agonists. The simulations reveal a transient intermediate state in the weak→strong transition.

Introduction

Earlier version: Skeletal muscle nicotinic acetylcholine receptors (AChRs) mediate synaptic transmission between motor neurons and muscle fibers. These pentameric ligand-gated ion channels (pLGICs) are composed of four different subunits (2 α1, β1, δ, and ε) arranged symmetrically around a central ion-conducting pore (Rahman et al., 2020; Unwin, 1993). Binding of the neurotransmitter acetylcholine (ACh) or other agonists to two target sites in the extracellular domain (ECD) increases the probability of a global conformational change that opens a distant (∼60Å) channel ‘gate’ in the transmembrane domain (TMD), to allow water and cations to cross the membrane.

Revised version: Adult vertebrate skeletal muscle nicotinic acetylcholine receptors (AChRs) are allosteric proteins that mediate synaptic transmission between motor neurons and muscle fibers. These pentameric ligand-gated ion channels are composed of four different subunits (2 α1, β1, δ, and ε) arranged symmetrically around a central ion-conducting pore (Rahman et al., 2020; Unwin, 1993). Binding of the neurotransmitter acetylcholine (ACh) or other agonists to two orthosteric sites in the extracellular domain (ECD) increases the probability of a global change in protein 3D conformation that activates a distant (~60.0 Å) allosteric site (a gate) in the transmembrane domain (TMD), allowing cations to cross the membrane.

Earlier version: At the start of the channel-opening isomerization the agonist binding sites switch from low to high affinity (LA→HA), and at the end of the isomerization the gate switches from closed to open (CLA→OHA) (Purohit et al., 2013). Without a bound agonist, the probability of the O shape (PO) is small and baseline activity is negligible (Jackson, 1986; Nayak et al., 2012; Purohit & Auerbach, 2009). Agonists (A) increase PO because they provide more favorable binding energy to the target in its active vs resting conformation, making OHA more stable than CLA. Agonists also bind strongly to the transition state that separates O from C and so also increase the channel-opening rate constant (Grosman et al., 2000).

Revised version: At the start of the channel-opening isomerization the neurotransmitter sites switch from having a low to a high affinity for agonists (LA→HA), and at the end the gate switches from closed to open (C→O) (Grosman et al., 2000; Purohit et al., 2013). Hence, the reaction scheme for the ‘gating’ allosteric transition is CLA⇄OHA. Without any bound agonists the probability of having the OHA shape (PO) is ~10−6 and baseline current is negligible (Jackson, 1986; Nayak et al., 2012; Purohit & Auerbach, 2009). With or without agonists, the receptor turns on and off spontaneously (influenced only by temperature). Agonists increase PO (that is, preferentially stabilize OHA) because they provide net favorable binding energy to their target when it is in the active, high-affinity conformation. Agonists also stabilize the transition state that separates CLA from OHA and therefore increase PO almost exclusively by increasing the channel-opening rate constant.

Added in the revised version: The difference in binding free energies, HA minus LA (ΔGHA-ΔGLA), drives the receptor’s allosteric transition and determines relative agonist efficacy (maximum PO). Another ligand attribute we use is efficiency that is a function of the ratio of binding free energies, (1-ΔGLA/ΔGHA) (Nayak et al., 2019). Efficacy (efficiency) is the amount (fraction) of binding free energy delivered to the gating apparatus. In a dose-response curve, efficacy is manifest as the high-concentration asymptote and efficiency as the correlation between this maximum and EC50.

Removed in the revised version: The binding free energy difference between high and low affinity (ΔGHA-ΔGLA) drives receptor activation and determines agonist efficacy, the high-concentration limit of the dose-response curve. In contrast, the binding free energy ratio (ΔGLA/ΔGLA) generates a correlation between the curve’s maximum and midpoint and determine agonist efficiency, the fraction binding energy converted into movements that trigger receptor conformational change.

Earlier version: Adult-type muscle AChRs have an unfavorable unliganded isomerization free energy (+8.3 kcal/mol at −100 mV) (Nayak et al., 2012) and 2 approximately equal and independent binding sites a (Nayak & Auerbach, 2017). For ACh, ΔGHA=-10.2 kcal/mol and ΔGHA =-5.1 kcal/mol per site, so adding their differences at 2 sites to that of the intrinsic isomerization, gives −1.9 kcal/mol. A free energy difference is proportional to the logarithm of an equilibrium constant. Hence, after binding 2 neurotransmitters the AChR gating equilibrium constant increases from 7.4×10−7 to 25, and PO increases from ∼10−6 to 0.96 (Fig. 1A). Binding energy from 2 neurotransmitters turns the receptor approximately from off to on.

Revised version: Adult-type muscle AChRs have 2 approximately equal and independent neurotransmitter sites (Nayak & Auerbach, 2017). The net LA→HA binding free energy change (ΔG) for two neurotransmitters (−10.2 kcal/mol) adds to that for the unfavorable unliganded isomerization (+8.3 kcal/mol at −100 mV) (Nayak et al., 2012) making the free energy change of the allosteric transition with 2 bound ACh favorable (−1.9 kcal/mol). ΔG is proportional to the logarithm of the corresponding equilibrium constant (by −RT, where R is the gas constant and T the absolute temperature, equal to −0.59 at 23 °C). Accordingly, after 2 neurotransmitters have bound, LA→HA rearrangements at the 2 orthosteric sites increases P from ~10−6 to ~0.96 (Figure 1A).

Earlier version: AChR activation (function and structure)

Revised version: AChR activation.

Earlier version: A. cyclic activation scheme. Vertical, gating. Receptors isomerize between resting (C) and active (O) conformations (equilibrium constant Ln, where n is the number of bound ligands). Horizontal, binding. Agonists (A) bind weakly to C (equilibrium dissociation constant KdC, free energy change ΔGLA) and strongly to O (KdO, ΔGLA). In adult AChRs the 2 sites are equivalent and independent and there is no significant external energy L2/L0 =(KdC/KdO)2 (Nayak & Auerbach, 2017). B. intermediate steps in binding and gating. Left, the agonist diffuses to the target to form an encounter complex that then becomes a LA complex by a ‘catch’ rearrangement. Right, The LA site becomes HA by a ‘hold’ local rearrangement, followed by domain restructurings that eventually generate a conducting pore. Bold, the. free energy changes in catch and hold (ΔGLA and ΔGHA) determine agonist action. C. overlay of α and δ subunits extracellular domains, toxin-less (red, 6UWX.pdb) and apo (blue, 7QKO.pdb). There are no major deviations (Cα RMSD = 0.3 Å). D. α-δ subunit neurotransmitter site with CCh (blue) in a HA, desensitized conformation (7QL6.pdb; (Zarkadas et al., 2022)). The agonist is in the trans orientation (tail points towards δ).

Revised version: A. Cyclic, thermalized activation scheme (vertical, gating; horizontal, binding; red, main physiological pathway). Receptors isomerize globally between closed-channel/low-affinity (CLA) and open-channel/high-affinity (OHA) conformations (equilibrium constant Ln; n, number of bound ligands). Agonists (A) bind weakly to CLA (equilibrium association constant KLA, free energy change ΔGLA) and strongly to OHA (KHA, ΔGHA). In adult AChRs the 2 orthosteric sites are approximately equivalent and independent, and there is no significant external energy: L2/L0 =(KHA/KLA)2 (Nayak & Auerbach, 2017). B. Top, binding and gating. Bottom, catch and hold. The agonist diffuses to the target to form an encounter complex (A-C), a local ‘catch’ rearrangement establishes the LA complex (ACLA), a local ‘hold’ rearrangement establishes the HA complex (ACHA), the rest of the protein isomerizes without a further change in affinity to generate a conducting pore (AO). Gray, steps that incur the same energy change for all agonists used in this study; black, agonist-dependent free energy changes in catch is ΔGLA and in hold (ΔGHA-ΔGLA). C. α-δ subunit extracellular domains, red, after toxin removal (6UWZ.pdb) and blue, apo (7QKO.pdb). There are no major deviations (Cα RMSD = 0.3 Å). D. Closeup of the desensitized Torpedo α-δ subunit neurotransmitter site occupied by carbamylcholine (CCh, blue) (7QL6.pdb; (Zarkadas et al., 2022)). In this is HA conformation, 3 aromatic groups in the α subunit (149-190-198) surround the agonist’s cationic center (+) together provide ~90% of the binding energy for ACh (Purohit et al., 2014), and the tail points toward he δ subunit (trans orientation).

Added in the revised version: The relationship between ΔGLA and ΔGHA binding free energies and the agonist properties affinity, efficacy and efficiency are considered elsewhere (Auerbach, 2024). Here, our goal was to associate the ΔGLA→ΔGHA change in binding free energy (4 agonists) with restructuring events at a neurotransmitter site. Molecular dynamics (MD) simulations were used to explore conformational changes in the LA→HA transition at the α−δ neurotransmitter site of the Torpedo AChR (6UVW; (Rahman et al., 2020)). In brief, binding free energies calculated in silico matched those estimated by using electrophysiology in vitro (Indurthi & Auerbach, 2023), allowing the identification of LA versus HA structures in the simulations. In the LA→HA transition for all agonists, three prominent rearrangements were rotation (flip) of the ligand, downward displacement (flop) of loop C, and compaction and stabilization of the binding cavity (fix). The initial LA and final HA conformations of the site were connected by a brief intermediate state. The simulations suggest that the movement of a weakly-bound agonist starts the change to strong binding and initiates the protein’s allosteric transition.

Removed in the revised version: Our goal is to associate ΔGLA→ΔGHA changes with underlying structural rearrangements at the binding site. Here, we use MD simulations to explore the conformational dynamics of the LA→HA transition at the α-δ neurotransmitter binding site of the Torpedo AChR (6UVW; (Rahman et al., 2020)). Binding free energies calculated in silico aligned with those estimated experimentally by using electrophysiology, so it was possible to interpret the agonist energy changes with restructuring of the binding pocket. It is difficult to draw sweeping conclusions because we examined only 4 agonists and only one wild-type site. However, in the LA→HA transition for all constructs the agonist somersaults, loop C droops to reposition αY190 and rearrange an aromatic pocket that contracts and becomes more hydrophobic. Also, in all systems the initial LA and final HA conformations of the binding site were connected by a brief intermediate state. For some but not all agonists, loop F moves in and a conserved salt bridge shortens. The results suggest that that agonist mobility is important determinant of the increase in binding free energy, and that rather than acting as a lid that closes, loop C is a positioning device that inserts an affinity-increasing side chain into the pocket core.

Results

Hold

Added in the revised version: Traditionally, receptor activation is divided into distinct stages called ‘binding’ (formation of a ligand-protein complex) and ‘gating’ (receptor isomerization) (Figure 1). Here, we are concerned with the connection between these stages as molecular movements at a neurotransmitter site.

Added in the revised version: In AChRs, agonist binding is a 2-step process. First, the agonist (A) diffuses to the target and forms an encounter complex (A+C⇄A-C), after which a local ‘catch’ rearrangement establishes the LA complex (A-C⇄ACLA) (Jadey & Auerbach, 2012). Gating, too, involves multiple steps (Gupta et al., 2017; Purohit et al., 2013). First, the site undergoes a local ‘hold’ rearrangement that establishes the HA complex (ACLA⇄ACHA), after which other protein domains restructure without changing agonist affinity, eventually opening the gate (Figure 1B).

Added in the revised version: The only agonist-dependent free energy changes in activation occur in the two local site rearrangements, ΔGLA (in catch) and ΔGHA−ΔGLA (in hold). The terminal state in the ‘binding’ stage of activation is ACLA and the initial state in the ‘gating’ stage is ACHA. Our interest here is hold, ACLA⇄ACHA, the local rearrangement that connects these two states. The free energy changes associated with the other events in the activation sequence, namely the formation of the encounter complex and downstream domain restructurings, are approximately the same for all agonists and so are not relevant to this study.

Removed in the revised version: Traditionally, receptor activation is divided into binding (formation of a ligand-protein complex) and gating (receptor isomerization). Here, we are concerned only with movements at binding sites that connect these two steps. In binding, the agonist (A) diffuses to interact with the target (A+C=AC), after which a local rearrangement establishes the LA complex (AC=ACLA) (Jadey & Auerbach, 2012). In gating, the site again rearranges to produce the HA complex (ACLA=ACHA), after which the rest of the protein domains (ECD, TMD and gate) restructure to establish AOHA. We call the two local rearrangements of the binding site ‘catch’ (free energy change ΔGLA) and ‘hold’ (ΔGHA-ΔGLA) (Fig. 1B). The end of binding is ACLA and the start go gating is ACHA, so hold is the local rearrangement that connects these two activation steps.

Earlier version: The free energy changes associated with the other events in binding and gating, namely diffusion to the site (the chemical potential) and restructuring of receptor domains, are the same for the agonists we have examined and so are not relevant to this study. Furthermore, energy coupling between the binding pocket and distant residues (>15 Å) typically is weak (Gupta et al., 2017) (Fig. S1A). Only catch and hold determine differences between agonists with regard to affinity (ΔGLA, ΔGHA), efficacy (ΔGHA-ΔGLA) and efficiency (1-ΔGLA/ΔGHA) (Nayak et al., 2019). In AChRs, ΔGLA and ΔGHA have been measured experimentally for 23 agonists and 53 binding site mutations (Indurthi & Auerbach, 2023) (Fig. S1B). The structures generated in the MD simulations were identified as ACLA or ACHA by aligning ΔGLA and ΔGHA values calculated in silico with those from experiments.

Revised version: Activation free energy changes incurred outside catch and hold, namely from the chemical potential associated with encounter complex formation (Phillips, 2020) and from rearrangements in other domains (and water), together represent an energy offset that is approximately agonist independent. With a few exceptions, energy coupling between the agonist binding pocket and distant residues (>15.0 Å) typically is weak (Gupta et al., 2017), so only the agonist’s energy changes in catch and hold determine the pharmacology attributes of affinity (ΔGLA or ΔGHA), efficacy (ΔGHA−ΔGLA) and efficiency (1-ΔGLA/ΔGHA). In AChRs, ΔGLA and ΔGHA have been measured experimentally by using electrophysiology for 23 agonists and 53 neurotransmitter site mutations (Indurthi & Auerbach, 2023).

Earlier version: The cryo-EM structures of AChRs include apo, resting-C with a bound toxin and high-affinity desensitized (Rahman et al., 2022; Rahman et al., 2020; Zarkadas et al., 2022). Binding sites of desensitized AChRs have a similar (perhaps identical) high affinity as do O sites (Auerbach, 2020). The end states of the hold rearrangement are short-lived and have not been captured as structures.

Revised version: Apo, resting-C with a bound toxin, and high-affinity desensitized AChR structures have been solved by using cryo-EM (Rahman et al., 2022; Rahman et al., 2020; Zarkadas et al., 2022). Desensitized AChRs have a similar (perhaps identical) high affinity as AOHA (Auerbach, 2020; Nayak & Auerbach, 2017). Below, we describe the conformation dynamics that connect the end states of the hold rearrangement, ACLA→ACHA (Figure 1B, bottom), neither of which has been captured as a structure. However, as described below, they can be identified in MD simulations that use resting-C as a start.

Removed in the revised version: Because ΔGLA and ΔGHA are independent of energy changes outside the binding pocket, in the simulations we removed the TMD and studied only α-Δ subunit ECD dimers. The starting structure for the hold transition was (PDB: 6UWZ) after toxin removal that is essentially the same as the apo structure (Zarkadas et al., 2022)) (Fig. 1C). The 4 agonists we examined were acetylcholine (ACh), carbamylcholine (CCh), epibatidine (Ebt) and epiboxidine (Ebx).

Docking

Earlier version: Binding energy changes in hold, ACLA=ACHA (Fig 1B). To understand the structural underpinnings, we began by removing α-bungarotoxin from the high-resolution (2.69 Å) structure 6UWZ (Rahman et al., 2020) and docking agonists into the empty pocket at the α-δ subunit interface (Fig. 1C). The docking simulations had strong binding scores, ranging from −5.82 to −9.18 kcal/mol, suggesting stable interactions. However, unlike the agonist orientation apparent in cryo-EM structures of desensitized AChRs (Gharpure et al., 2019; Noviello et al., 2021; Rahman et al., 2022; Walsh et al., 2018; Zarkadas et al., 2022) (Fig. 1D), the agonist’s tail pointed away from (cis) rather than towards (trans) the complimentary (non-α) subunit. This was the case for the top 3 docking scores for all 4 ligands (Fig. 2A). Because the cis orientation was not observed in any cryo-EM (HA) structure, we examined the top 200 docking scores for CCh. All poses were cis, with the tail pointing towards the α subunit. The cis orientation was also apparent in docking of agonist into a putatively LA binding site of a homology model (Tripathy et al., 2019). To start, we associated the agonist orientation in the hold end states as cis in ACLA versus trans in ACHA.

Revised version: We began by removing α-bungarotoxin from the high-resolution (2.69 Å) structure 6UWZ (Rahman et al., 2020) (Figure 1C) and docking agonists into the empty pocket. The docking simulations had strong binding scores ranging from −5.82 to −9.18 kcal/mol (Figure 2A), suggesting that the interactions were stable. However, unlike the agonist orientation apparent in cryo-EM structures of desensitized (HA) AChRs (Noviello et al., 2021; Rahman et al., 2022; Rahman et al., 2020; Walsh et al., 2018; Zarkadas et al., 2022) (Figure 1D), the docked agonist’s tail pointed towards (cis) rather than away (trans) from the α subunit (Figure 2B, top). This was the case for the top 3 docking positions for all 4 ligands.

Earlier version: Agonist docking and dynamics

Revised version: Agonist docking and loop dynamics.

Earlier version: A. top, Agonist structures (blue, main N+): carbamylcholine (CCh), acetylcholine (ACh), epibatidine (Ebt) and epiboxidine (Ebx). B. top, root-mean-square-deviation (RMSD) of the system Cα (mean±SD, triplicates) throughout 200 ns MD simulations (ACh, cyan; CCh, green; Ebt, orange; Ebx, purple). bottom, close-up view of the pocket showing movements during MD simulations with CCh (red, apo; orange, 0 ns; blue, 200 ns;). Loop C moves down, loop F closes in, the agonist rotates cis→trans. C. α-Δ binding pocket (6UVW.pdb minus toxin) with agonists (top 3 poses). top, docked into the 0 ns structure (red). Loop C is up, loop F is out, agonist is cis. bottom, docked into the 200 ns structure after removing CCh (blue). Loop C is down, loop F is in and the agonist is trans. D. Docking score energies (mean±SD), 0 ns (red) and 200 ns (blue) suggest a LA→HA transition during the course of the simulation.

Revised version: A. Top, agonists (blue, cationic center): carbamylcholine (CCh), acetylcholine (ACh), epibatidine (Ebt) and epiboxidine (Ebx). Bottom, α−δ site (6UVW.pdb minus toxin, resting-C) with docked agonists (top 3 poses). docked into resting-C (red): loop C is up, agonist is cis. docked after 200 ns simulation and removing CCh (blue): loop C is down, agonist is trans. B. Bottom, docking scores (mean±SD, n=3). red, resting-C; blue, 200 ns. C. Cα RMSD (mean+SD, triplicates) (ACh, cyan; CCh, green; Ebt, orange; Ebx, purple). D. close-up of the CCh-occupied pocket. Red, resting-C; orange, equilibrated (0 ns); blue, 200 ns. After 200 ns loop C has flopped, loop F has moved inward, and the agonist has flipped cis→trans.

Added in the revised version: Because the cis orientation was not observed in any cryo-EM structure, we examined the top 200 docking scores for CCh. All poses were cis, with the tail pointing towards the α subunit and away from the main binding cavity. Apparently, in the apo structure the trans orientation is energetically unfavorable. Hence, to start our analysis of the hold rearrangement we hypothesized that the agonist orientations in the end states are cis in ACLA versus trans in ACHA.

Earlier version: LA→HA transition

Revised version: The LA→HA transition

Earlier version: We performed 200 ns MD simulations for each of the docked complexes, in triplicate (with different seed values) (Fig. 2B top). The MD simulations for all three replicates demonstrated similar root-mean-square-deviation (RMSD) patterns, with differences consistently below 1Å (Table S1) suggesting that the trajectories were robust and reproducible. Although there were initial deviations, all systems became stable by ∼120 ns. The root-mean-square-fluctuation (RMSF) of the Cα atoms, employed to identify regions of flexibility, showed instability mainly in loops C and F around the pocket (Fig. S2). The fluctuations observed at the base of the ECD were anticipated because the TMD that offers stability here was absent.

Revised version: We performed 200 ns MD simulations for each of the docked complexes, in triplicate (with different seed values). Video 1 shows a simulation with CCh on 2 times scales, expanded for the initial 200 ns and condensed for the full 1.0 μs extension. The 200 ns simulations for all three replicates for all agonists showed similar root-mean-square-deviation (RMSD) patterns with differences consistently below 1Å (Figure 2C, Figure 2-Source Data 1) suggesting that the trajectories were sufficiently robust and reproducible. Although there were initial deviations, all systems became stable by ~120 ns. The root-mean-square-fluctuation (RMSF) of the Cα atoms showed flexibility around the pocket mainly in loops C and F (but not E) (Figure 2-figure supplement 1). The fluctuations at the base of the ECD were anticipated because the TMD that offers stability here was absent. As described below, the match between experimental binding energies (whole receptors) and those calculated in the simulations (ECD dimers) is further evidence that the ECD-TMD interfacial region does not influence agonist binding energy.

Added in the revised version: The ECD dimer in the absence of a bound agonist (apo) exhibited a similar RMSF and RMSD pattern for a 200 ns simulation, to that observed with ligand present (CCh) (Figure 2-figure supplement 2A, Figure 2-figure supplement 2B). The displacement of loop C apparent with agonists (see below) was also present but less pronounced in apo (Figure 2-figure supplement 2C).

Earlier version: In addition to the fluctuations in these loops, during the 200 ns simulations the ligand orientation changed dramatically, cis-to-trans. At the start of the simulations the agonist’s tail points towards α, and at the end it points towards δ (Fig. 2B bottom). This rearrangement represents a substantial rotation (a ‘flip’) about the agonist’s main cationic center (N+) that remains close to αW149 throughout the somersault. The final MD configuration (with CCh) was well-aligned with the CCh-bound cryo-EM desensitized structure (7QL6.PDB) (RMSD < 0.5 Å) (Fig. S3), further demonstrating that the simulation had converged.

Revised version: In addition to the fluctuations in loops C and F during the simulations, the ligand orientation inverted, cis→trans. For all ligands, at the start of the simulations the agonist’s tail points towards α, but by the end points towards δ (Figure 2D). The agonist re-orientation represents a nearly complete somersault (a ‘flip’) about the agonist’s main cationic center (N+) that remains close to αW149 throughout (see Figure 5). The final MD configuration (with CCh) aligns well with the CCh-bound cryo-EM desensitized structure (7QL6.PDB) (RMSD < 0.5 Å) (Figure 2-figure supplement 3). Video 1 shows that the final trans orientation with CCh remains constant for at least 1 μs, indicating that structure had settled by 200 ns, with critical rearrangements conserved.

Earlier version: To test whether rearrangements at the binding pocket in the simulations themselves enforce the trans configuration, we removed the bound CCh from the final state MD structure and re-docked all four agonists. The preferred poses for all 4 agonists were all trans (Fig. 2C). In addition, binding energies from docking scores were higher for the converged MD structure compared to the initial binding site structure. As a further control we carried out MD simulations of a pentamer docked with ACh and found similar structural changes at the binding pocket compared to the dimer (Fig. S4). These results indicate that a dimer is appropriate to study local rearrangements of an AChR neurotransmitter binding site and suggest that in the MD simulations the binding site restructures from ACLA to ACHA conformation (Fig. 2D).

Revised version: To test whether rearrangements at the binding pocket in the simulations enforced the trans configuration, we removed the bound CCh from the final 200 ns MD structure and re-docked the agonists. The preferred poses for all 4 ligands were trans (Figure 2A, bottom). In addition, binding energies from docking scores were higher for the 200 ns structures compared to the initial binding site structures, consistent with a LA-to-HA transition (Figure 2B). As a further control we carried out MD simulations of a pentamer docked with ACh and found similar structural changes at the binding pocket compared to the dimer (Figure 2-figure supplement 4).

Added in the revised version: Together, these results suggest that ECD dimers are adequate to study local rearrangements of the α–δ AChR neurotransmitter binding site. Further, they provide support for the hypotheses that in the simulations the site restructures from ACLA to ACHA, and that the preferred agonist pose is cis in LA versus trans in HA structures.

Removed in the revised version: Principal component analysis

Removed in the revised version: To identify the most prominent conformational states of the receptor during the MD simulations, principal component analysis (PCA) was carried out based on the most significant fluctuations (principal components). Considering the first two principal components (PC-1 and PC-2) that capture the most pronounced Cα displacements, the protein backbone fluctuations revealed three energy minima (m1, m2, and m3) that mainly reflect the stabilities of loops C and F (Fig. 3). For all four agonists, the trajectories quickly reach m1 (≤20 ns), then transition through m2, and culminate in m3 (180-200 ns). In the simulations progress in time was sequential, m1→m2→m3.

Principal component analysis (PCA)

Added in the revised version: To identify the most prominent conformational states of the receptor during the simulations, PCA was carried out based on the most significant fluctuations (principal components). Considering the first two principal components (PC-1 and PC-2) that captured the most pronounced Cα displacements (Figure 3-Source Data 1), backbone fluctuations revealed three energy minima (m1, m2, and m3) that reflect the stabilities of loops C and F (Figure 3). To assess consistency across independent simulations, we calculated the inner products of the top 10 principal components from each run. The resulting heatmaps reveal the degree of similarity between runs, with the inner product values of PC-1 and PC-2 reflecting considerable overlap (50-90%) in the principal dynamic modes captured in each simulation (Figure 3-figure supplement 1).

Added in the revised version: Principal Component Analysis (PCA).

Earlier version: For each agonist, the left panel plots PC1 vs PC2, the first two principal components that capture the maximum variance in the trajectory. Colors represent free energy value in kcal/mol (upper left, bottom). For all agonists there are 3 energy minima (red), m1, m2, and m3, that correspond to different conformations of the protein-ligand complex. The right panels are ‘porcupine’ plots that indicate the direction and magnitude of changes between PC1 and PC2. In order of appearance, m1 (≤20 ns)< m2<m3 (≥180 ns) (see text). m1 represents ACLA, m3 represents ACHA, and m2 is an intermediate configuration (corresponding ΔGs given in Table 1).

Revised version: For each agonist, the left panel plots PC1 vs PC2, the first two principal components that capture the maximum variance in the trajectory (Figure 3-Source Data 1). Colors represent free energy value in kcal/mol (scale, upper left-bottom). For all agonists there are 3 energy minima (dark red), m1, m2, and m3, that correspond to different conformations of the protein. The right panels are ‘porcupine’ plots indicating that the direction and magnitude of changes between PC1 and PC2 is in loops C and F. From energy comparisons (Figure 4, Figure 4-Source Data 1) and temporal sequences (Figure 3-figure supplement 2, Figure 4-Source Data 1): m1 is ACLA, m3 is ACHA, and m2 is an intermediate configuration.

Earlier version: Accordingly, we hypothesized that m1 corresponds to ACLA and m3 corresponds to ACHA. For all agonists, the m2 population has greater variance (more instability) compared to m1 and m3. This, along with its position in the sequence, suggests that m2 represents an unstable intermediate configuration within hold, between ACLA and ACHA. For all agonists and trajectories, m3 had the least variance (was most stable), again supporting convergence by 200 ns.

Revised version: For all 4 agonists, trajectories started in m1 and ended in m3 (Figure 3-figure supplement 2). Accordingly, we hypothesized that m1 corresponds to ACLA and that m3 corresponds to ACHA. For all agonists, the m2 population had a greater variance (was less stable) compared to m1 and m3. This, along with its temporal position, suggests that m2 represents an unstable intermediate configuration within the hold transition, between ACLA and ACHA. For all agonists and trajectories, m3 had the least variance (was the most stable).

Added in the revised version: The pronounced fluctuations in loop C are evident in residue-wise RMSF plots and align well with the contributions of the first two principal components (PC-1 and PC-2) (Figure 2-figure supplement 2D). The apo and CCh occupied proteins adopt similar conformational spaces during the 200 ns simulation (Figure 2-figure supplement 2E). However, presence of the agonist appears to increase conformational diversity, generating a broader distribution in the PCA plot. The bottom of the m1, m2 and m3 energy wells in Figure 3 represent the most stable configurations of each population and serve as reference points.

Removed in the revised version: Transitions between the 3 energy minima varied with the agonist. Always, m1 was apparent first (at 5-14 ns for ACh; 15-24 ns for CCh; 21-41 ns for Ebt; 20-56 ns for Ebx), followed by m2 (20-50 ns; 50-60 ns; 60-100 ns; 100-120 ns). All systems converged to m3 (180-200 ns) (Fig. 3). The bottom of the m1, m2 and m3 energy wells represent the most stable configurations of each population and serve as a reference point for the rest of energy well.

Removed in the revised version: Although the bottom of the well for 3 energy minima from PCA represent the most stable overall conformation of the protein, they do not convey direct information regarding agonist stability or orientation. To incorporate this information, we performed a cluster analysis of the ligand using frames from the bottom of each PCA well as inputs. The top three clusters, each having RMSD of ≤ 1Å, shared a similar ligand orientation (Fig. S5) and were selected to compute binding free energy using MM-PBSA. The remaining frames were disregarded. The fraction of frames accepted for free energy calculations ranged from a minimum of 20% (ACh, m1) to a maximum of 71% (Ebt, m3). On average, the fraction of frames selected from each energy minimum was ∼50%, showing the dynamic nature of the ligand in the pocket (Table 1).

Earlier version: Computed vs experimental binding energies

Revised version: Computed versus experimental binding energies

Added in the revised version: Although the well bottoms in Figure 3 represent the most stable overall conformations, they do not directly convey information regarding agonist stability or orientation. To incorporate this information, we performed a cluster analysis of the ligand configuration using frames from the bottom of each PCA well as input. The top three clusters, each having RMSD of ≤ 1.0 Å, shared a similar ligand orientation (Figure 3-figure supplement 2) and were selected to compute binding free energies (Figure 4, Figure 4-Source Data 1, Figure 4-Source Data 2). The remaining frames were disregarded. The fraction of frames from each minimum accepted for free energy calculations ranged from 20% (ACh, m1) to 71% (Ebt, m3). On average, the fraction of frames selected was ~50%, showing the dynamic nature of the ligand in the pocket (Video 1).

Removed in the revised version: The above procedure defined populations of structures that we hypothesize represent the end states of hold, ACLA (m1) and ACHA (m3), plus an intermediate (m2). We used MM-PBSA (Molecular Mechanics Poisson-Boltzmann Surface Area) computations of the selected structures to estimate binding energy of each agonist to each population. MM-PBSA is particularly useful to compare binding affinities between different states and among agonists (Kollman et al., 2000; Rastelli et al., 2010; Sun et al., 2014). In our use, MM-PBSA computations included both enthalpic (ΔH) contributions from molecular interactions (van der Waals, electrostatics, solvation) and entropic (TΔS) contributions calculated using normal mode analysis. By capturing the dynamic interplay between the ligand and receptor as well as the influence of the surrounding solvent environment this approach enables computation of a Gibbs free energy (ΔG). However, there are limitations with regard to MM-PBSA accurately predicting absolute binding free energies (Genheden & Ryde, 2015; Hou et al., 2011) that depends on parameterization of the ligand (Oostenbrink et al., 2004).

Removed in the revised version: The calculated and experimental ΔG values are shown in Table 1. For ACh and CCh, there is good agreement between ΔGm1 and ΔGLA and between ΔGm3 and ΔGHA. For these agonists, in silico values overestimated experimental ones only by ∼8% and ∼25%. The agreement was not as good for the other 2 agonists, as calculated values overestimated experimental ones by ∼45% (Ebt) and ∼130% (Ebt). However, the fractional overestimation was approximately the same for ΔGLA and ΔGHA.

Removed in the revised version: Despite these discrepancies in absolute values, the trend of an increase in binding free energy, m1 versus m3, is consistent with the assignment, ACLA versus ACHA. For all agonists, the stability of the ligand-receptor complex increases significantly in the m1→m3 transition (Fig. 4A). Further support for the assignment is from comparing calculated vs experimental efficiencies (η) that depend on ΔGLA/ΔGHA (Nayak et al., 2019). Table 2 shows that η values calculated in silico agree closely (≤5%) with those measured experimentally, for all agonists. Apparently, MD simulations of a single pdb file successfully model the hold transition. The match between calculated and experimental efficiencies is evidence that that m1 represents low-affinity ACLA and m3 represents high-affinity ACHA.

Binding free energies and pocket properties.

Added in the revised version: A. Calculated (yellow) versus experimental (blue) binding free energies for 4 agonists (Figure 2A, top) (Figure 4-Source Data 1). PBSA calculations were done on clusters selected from m1 and m3 minima of PCA plots (Figure 3; Figure 3-figure supplement 2). Left, absolute ΔG and right, efficiency. The agreement in efficiencies supports the hypothesis that m1 represents ACLA and m3 represents ACHA. B. In ACLA→ACHA, VdW interactions increase (left), pocket volume decreases (center) and the number of water molecules decreases (right). Overall, the pocket stabilizes, compacts and de-wets.

Added in the revised version: The above procedure defined populations of structures that we hypothesized represent the end states of hold, ACLA (m1) and ACHA (m3), plus a previously unidentified intermediate state (m2). We used PBSA (Poisson-Boltzmann Surface Area) computations of the selected structures for each agonist in each cluster population and compared binding free energies calculated in silico with those measured previously in vitro (Indurthi & Auerbach, 2023). There are limitations with regard to PBSA accurately predicting absolute binding free energies (Genheden & Ryde, 2015; Hou et al., 2011) that also depend on parameterization of the ligand (Oostenbrink et al., 2004). However, because we knew the actual binding free energies from electrophysiology experiments we used the PBSA calculations only to test our hypothesis that m1 is ACLA and m3 is ACHA.

Added in the revised version: Calculated PBSA and experimental ΔG values are compared in Figure 4A left (Figure 4-Source Data 1). For ACh and CCh there was good agreement between ΔGm1 and ΔGLA and between ΔGm3 and ΔGHA. For these agonists, in silico values overestimated experimental ones by ~8% and ~25%. The agreement was not good for the other 2 agonists, with calculated values overestimating experimental ones by ~45% (Ebt) and ~130% (Ebt). However, in all cases ΔGm3 was more favorable than ΔGm1, and the % overestimation was approximately the same for m1 versus m3. Despite the mismatches with the PBSA calculations, the comparison is consistent with the hypothesis.

Added in the revised version: Further support for the assignments comes from comparing calculated vs experimental efficiencies that depend on the binding free energy ratio, ΔGLA/ΔGHA, rather than absolute ΔGs. As described elsewhere, efficiency is the agonist’s free energy change in hold relative to its total free energy change, (ΔGHA-ΔGLA)/ΔGHA or (1 −ΔGLA/ΔGHA) (Nayak et al., 2019). Figure 4A right (Figure 4-Source Data 1) shows that efficiency values calculated in silico agree almost perfectly with those measured experimentally. This result not only further supports the hypothesis that m1 represents ACLA and m3 represents ACHA, but also indicates that efficiency can be computed from structure alone. We conclude that MD simulations of a single pdb file successfully model the end states of the hold transition and that the trajectories likely reproduce the conformational dynamics in ACLA→ACHA.

Removed in the revised version: A. Calculated Gibbs free energy (ΔG) using MM-PBSA for the three distinct energy minima (m1, m2, and m3) showing the progression LA to HA (agonists, Fig. 2A). ΔGm2 is intermediate for ACh and CCh but least stable for Ebt and Ebx. B. vdW interactions increase m1 (ACLA) to m3 (ACHA), consistent with increased hydrophobicity. C. Binding pocket volume reduces m1 to m3, reflecting the more tightly packed pocket in ACHA. D. The number of water molecules in the pocket decreases m1→m3, orange, CLA; blue, CHA, reflecting a more hydrophobic environment in ACHA.

Earlier version: In the ACLA→ACHA conformational change, all agonists show an increase in van der Waals (vdW) interaction energy. The increase was most pronounced with Ebt (−9.77 kcal/mol) and least for Ebx (−3.54 kcal/mol), with CCh (−8.81 kcal/mol) and ACh (−6.71 kcal/mol) falling in between (Fig. 4, Table S2). Also, in ACLA→ACHA, restructuring of the binding pocket further contribute to generating a smaller, more-tightly packed pocket. This is shown in Fig. 4C as a reduction in binding pocket volume and in Fig. 4D as a decrease in the number of water molecules. The agonist’s flip into the pocket results in de-wetting and a more hydrophobic environment.

Revised version: All agonists show an increase in van der Waals (VdW) interaction energy in this rearrangement (Figure 4B). The calculated increase (PBSA calculations) was most pronounced with Ebt and least for Ebx, with CCh and ACh falling in between (Figure 4-Source Data 2). ACLA→ACHA restructuring of the binding pocket generates a smaller, more-tightly packed pocket (Tripathy et al., 2019), in part because of the increased VdW interactions (Video 1). This is shown in Figure 4C as a reduction in binding pocket volume and in Figure 4D as a decrease in the number of water molecules. The VdW contacts that form concomitant to agonist flip and loop C flop establish a more compact, hydrophobic and stable local environment (Video 1).

Earlier version: Flip-Flop

Revised version: Flip-Flop-Fix

Added in the revised version: In ACLA→ACHA, a cis→trans rotation of the agonist (flip) was prominent in all simulations (Figure 5). All agonists underwent a >130° rotation about the N+-αW149 fulcrum. Also, in all instances the flip occurred near the start of the hold rearrangement.

Removed in the revised version: The ligand somersault (flip) was prominent in all simulations of hold. All agonists undergo a large rotation about the N+-αW149 fulcrum (Fig. 5; Video 1). The extent of rotation was 131° for CCh, 156° for ACh, 148° for Ebt and 176° for Ebx. Another consistent and prominent restructuring event in ACLA→ACHA was the downward displacement (flop) of loop C toward the core of the pocket. This movement is conserved in related receptors and thought to play a role in receptor activation by agonists (Basak et al., 2020; Hansen et al., 2005; Hibbs et al., 2009) (see below). In the simulations, we quantified loop C flop by measuring the displacement of its tip (Cα of αC192) relative to its starting position (at 0 ns). The distance traveled at 200 ns varied with the agonist and was 9.2 Å for CCh, 5.5 Å for Ebt, 5.2 Å for ACh and 3.3 Å for Ebx (Fig. 5). We do not discern a pattern with regard to agonist affinity, efficacy or efficiency.

Agonist and loop movements in hold.

Earlier version: In each panel, left, superimposed cartoons of ACLA (m1; orange) and ACHA (m3; blue). Loop C is upper left and loop F is lower right. In ACLA→ACHA (orange→blue) there is a cis→trans reorientation of the agonist (a flip) and a downward movement of loop C (a flop). right, the degree of flip from m1 (red) to m2 (yellow) to m3 (blue). The range, m1 (ACLA) →m3 (ACHA), is 131° to 176°.

Revised version: In each panel left, superimposed cartoons of ACLA (m1; orange) and ACHA (m3; blue). Loop C is upper left and loop F is lower right. In ACLA→ACHA (red→blue) there is a cis→trans reorientation of the agonist (a flip) and a downward movement of loop C (a flop, orange arrow); right, flip from m1 (red) to m2 (yellow) to m3 (blue); degree, m1→m3.

Added in the revised version: A second consistent and prominent restructuring event in ACLA→ACHA was the downward displacement (flop) of loop C toward the pocket center (Cα of αW149). This well-known ‘clamshell closure’ motion is conserved in related binding proteins and thought to play a role in AChR activation by agonists (Basak et al., 2020; Hansen et al., 2005; Hibbs et al., 2009). In the simulations, loop C flop was measured as the displacement of its tip (Cα of αC192) in m3 relative to its apo position. The distance traveled varied with the agonist (Figure 6-table supplement 1) but we do not discern a pattern in the extent of loop C flop with regard to agonist affinity, efficacy or efficiency.

Earlier version: In hold, residues that form the pocket cavity move, apparently to accommodate the re-orientation of the ligand. In ACLA αY190, αY93, and δW57 are spaced apart, providing an open slot that accommodates the agonist’s tail in the cis orientation. However, in ACHA αY190 is positioned closer to αY93, effectively creating a unified surface that disallows the cis position of the tail (Fig. S6). However, in the simulations the event sequence is not clear so perhaps flip precedes flop, and agonist re-orientation first creates the space for αY190 to occupy.

Revised version: In hold, residues that form the pocket cavity move, presumably to accommodate the reoriented ligand. In ACLA, αY190, αY93, and δW57 are spaced apart, providing an open slot that accommodates the agonist’s tail in the ACLA, cis orientation (Figure 5-figure supplement 1). However, in ACHA loop C flop repositions αY190 closer to αY93, effectively filling the gap to create a unified surface.

Added in the revised version: The αW149, αY190, αY198 and αY93 rings surround the N+ (Figure 1D). In adult-type mouse AChRs, Ala substitutions of these 4 aromatic residues have different consequences and result in ACh LA binding free energy losses of 3.0, 2.7, 2.2 and 1.1 kcal/mol, respectively (Purohit et al., 2014). δW57A does not change ACh binding energy (Nayak & Auerbach, 2013). αW149 (in loop B) makes a cation-π bond with N+ that is important for ligand stabilization (Xiu et al., 2009; Zhong et al., 1998). In the simulations, the distance between αW149 and N+ decreases only slightly, AC →AC (Figure 6, Figure 6-table supplement 1). This cation-π bond appears to act as a slippery anchor (a hinge) that maintains stability of the N+ position even as the tail rotates about it.

Removed in the revised version: αW149 (in loop B) makes a cation-π bond with N+ that is important for ligand stabilization (Xiu et al., 2009; Zhong et al., 1998). In the simulations, the position and orientation of αW149 relative to N+ remained almost unchanged for all 4 agonists, in both ACLA (cis) and ACHA (trans) configurations (Fig. 6). Apparently, the cation-π bond acts as a slippery anchor that maintains stability of the N+ position even as the tail rotates about it.

Earlier version: Representative snapshots in hold

Revised version: Representative snapshots in ACLA→ACHA (hold).

Earlier version: left, rearrangements of loop C, loop F and the ligand (m1, red; m2, yellow; m3, blue). Right, conserved residues and ligand orientation. m1 is ACLA and m3 is ACHA. In m1, a functional group at the agonist tail interact with αY93 (all agonists) and αD200 (CCh and Ebt). The position and orientation of αW149 relative to N+ of the agonist remains unchanged m1→m3 and serves as a fulcrum for the cis→trans flip (see Fig. 5). In m2, the functional nitrogen at the agonist tail (CCh, Ebt and Ebx) interacts with the hydroxyl group of αY198. For all ligands, loop C flop (m1→m3) repositions αY190 in pocket. In m3, the agonist fully flips to trans, facilitating the interaction of N+ with αY190 and the formation of water-mediated hydrogen bond with the reactive group at the tail with δN109/δL121 (loop E).

Revised version: left, rearrangements of loop C, loop F and the ligand (red, m1; yellow, m2; blue, m3); right, residue and ligand orientations. m1 is ACLA, m2 is an intermediate state, m3 is ACHA. In m1, a functional group in the agonist tail interacts with αY93 (all agonists) and αD200 (only CCh and Ebt). The position and orientation of αW149 relative to N+ of the agonist remains nearly unchanged m1→m2→m3 and serves as a fulcrum for the cis→trans flip (see Figure 5). In m2, the functional nitrogen at the agonist tail (CCh, Ebt and Ebx) interacts with the hydroxyl group of αY198. For all ligands, αY190 repositioning and loop C flop (m1→m3) are correlated. In m3, the agonist fully flips to trans, facilitating VdW interactions, de-wetting, and the formation of water-mediated hydrogen bonds with the reactive group at its tail with δN109/δL121 backbone (loop E) via a structural water.

Removed in the revised version: Loop C flop results in re-positioning of αY190 into the aromatic pocket. αY190 is the main ligand binding energy source, nearly twice as strong as αW149 (Purohit et al., 2014), perhaps because deprotonation of its hydroxyl group generates a cation-anion bond with N+ (Bruhova & Auerbach, 2017). An interaction between αY190 and the αK145-αD200 salt bridge, suggested to be the trigger for the gating isomerization (Bruhova & Auerbach, 2017; Mukhtasimova et al., 2005). In ACLA→ACHA, loop C flop restructures the pocket to put N+ between the αY190 and αW149 rings, and αY190 closer to the salt bridge (Fig. 6).

Removed in the revised version: Other hold rearrangements

Earlier version: In ACLA (m1), the cis orientation is stabilized by the hydroxyl of αY93 that is proximal to a functional group on the agonist tail. Also, the tail nitrogen in CCh, Ebt, and Ebx (amide or amine) forms a hydrogen bond with the hydroxyl group of αY198, but only in the intermediate m2 state. This interaction appears to guide the flip and provide transient stability to the intermediate state. The absence of a secondary nitrogen in ACh possibly explains its alternate re-orientation route, as only this ligand was observed to flip counterclockwise (Fig. 5). It may also explain why the m2 basin in the PCA plots for ACh is especially broad and shallow (Fig. 3).

Revised version: In AChBP (perhaps CLA), the hydroxyl group of αY93 forms a hydrogen bond with the ligand head group (Celie et al., 2004; Hansen et al., 2005). In the AChR ACLA structure, the cis orientation appears to be stabilized by this hydroxyl as it is close to a functional group on the agonist’s tail. Likewise, the tail nitrogen in CCh, Ebt, and Ebx (amide or amine) forms a hydrogen bond with the hydroxyl group of αY198, but only in the intermediate m2 state. This interaction may guide the flip and provide transient stability. The absence of a secondary nitrogen in ACh possibly explains its alternate re-orientation route, as only this ligand was observed to flip counterclockwise (Figure 5). It may also explain why the m2 basin in the PCA plots for ACh is especially broad and shallow (Figure 3).

Removed in the revised version: In ACHA (m3), the αY190/αY198/δW57 rings and the αW149 indole interact with N+, and the hydroxyl group of αY93 forms a hydrogen bond with the ligand head group (Celie et al., 2004; Hansen et al., 2005). However, Ala substitutions of these 5 aromatic residues have quite different consequences and result ACh binding free energy losses of +4.4, +1.9, +0.1, +2.3 and +0.9 kcal/mol, respectively (Purohit et al., 2014). In ACHA, the tail functional moiety (carbonyl or secondary ring amine) in the trans orientation is sandwiched between αC192, αC193 and αY198 (loop C), αT150 (loop B) and δC108, δN109, δL111, and δL121 (loop E) (Fig. S7).

Removed in the revised version: αK145 is complex and undergoes agonist-dependent hold rearrangements. In ACHA this residue i) makes a salt bridge with αD200, ii) has a cation-π interaction with αY93, iii) contacts αY190 (possibly through the hydroxyl group) and iv) approaches loop F of the complementary subunit (Zarkadas et al., 2022). Further, the mutation αK145A lowers the efficiencies of ACh and CCh (from 0.51 to 0.41) but has no effect on those of Ebt and Ebx (0.41) (Indurthi & Auerbach, 2023).

Removed in the revised version: This agonist dependence with regard to function has a parallel in structure. In the hold simulations, the αK145-αD200 distance (Å) increases from 4.4 in apo to 6.1 or 5.1 in ACLA (CCh or ACh). The spread may relate to the agonist tail being inserted into the αY93 slot in the cis orientation. In ACLA →ACHA, for some agonist the αK145-αD200 distance (Å) shortens substantially to 2.6 or 2.8, making the bridge more rigid (ACh or CCh). However, with Ebt and Ebx the αK145-αD200 distance remains unstable, fluctuating between 3-6 Å throughout the course of the simulation. Likewise, αK145-αY190 distance stabilizes at 5.3 Å or 6.1 Å in ACHA with ACh and CCh but continues to fluctuate with Ebt and Ebx. In contrast, the αY190-αD200 distance remains stable in ACHA for all ligands. This suggests that the instability in the salt αK145-αD200 and αK145-αY190 distances can primarily be attributed to αK145 (Fig. S8).

Earlier version: In ACHA, the functional group found in the tail of each agonist -- the carbonyl group in CCh and ACh and the secondary ring amine in Ebt and Ebx - forms a hydrogen bond with the carbonyl of δN109 and the amine of δL121 in the complementary δ subunit via a structural water (Fig. 6), as reported elsewhere (Blum et al., 2010; Gharpure et al., 2019; Zarkadas et al., 2022). This water-mediated bonding appears to be crucial for stabilizing the agonist in the trans orientation. The start of the agonist’s flip, m1→m2, begins the establishment of this water-mediated bond that is completed in m2→m3, with all 4 agonists (Fig. 6).

Revised version: The ACLA→ACHA transition compacts, dehydrates and stabilizes the binding cavity (Figure 4; Video 1). In ACHA, the tail functional moiety (carbonyl or secondary ring amine) in the trans orientation is sandwiched between αC192, αC193 and αY198 (loop C), αT150 (loop B) and δC108, δN109, δL111, and δL121 (loop E) (Figure 6-figure supplement 1). As reported elsewhere (Blum et al., 2010; Rahman et al., 2020; Zarkadas et al., 2022), in ACHA the functional group found in the tail of each agonist -- the carbonyl group in CCh and ACh and the secondary ring amine in Ebt and Ebx - forms a hydrogen bond with the carbonyl of δN109 and amine of δL121 in the complementary δ subunit backbone via a structural water (Figure 6). The agonist’s flip occurs at the start of m1→m2, and the establishment of this water-mediated H-bond completes in m2→m3 (Figure 6) (Video 1). This water-mediated bonding and the increase in VdW interactions appear to both be crucial for stabilizing the loops and agonist in the HA conformation of the neurotransmitter site.

Added in the revised version: Salt-bridge

Added in the revised version: Near the pocket, the αK145-αD200 salt bridge (Figure 1D) has been suggested to play an important role in receptor activation (Beene et al., 2002; Gay & Yakel, 2007; Mukhtasimova et al., 2005; Padgett et al., 2007; Pless et al., 2011). The hold rearrangements of αK145 are complex and agonist dependent. In ACHA, this residue i) makes an apparent contact with αY93, ii) contacts the αY190 hydroxyl group (Figure 6-table supplement 1) and, iii) approaches loop F of the complementary subunit.

Added in the revised version: Interestingly, the mutation αK145A lowers the efficiencies of ACh and CCh (from 0.51 to 0.41) but has no effect on those of Ebt and Ebx (0.41 and 0.46) (Indurthi & Auerbach, 2023). This agonist dependence in function has a parallel in structure. In the hold simulations, the αK145-αD200 distance (Å) increases from only 4.4 in apo to 6.1 or 5.1 in ACLA (CCh or ACh). This spreading may relate to the agonist’s tail being inserted into the αY93 slot in the cis orientation. In ACLA→ACHA, for some agonists the αK145-αD200 separation (Å) shortens substantially to 2.8 or 2.6, making it more rigid (ACh or CCh). However, with Ebt and Ebx the αK145-αD200 distance remains unstable, fluctuating between 3-6 Å throughout the course of the simulation. Likewise, αK145-αY190 distance stabilizes at 5.3 Å or 6.1 Å in ACHA with ACh and CCh but continues to fluctuate with Ebt and Ebx. In contrast, the αY190-αD200 distance remains stable in ACHA for all ligands. This suggests that the instability in the αK145-αD200 and αK145-αY190 distances can primarily be attributed to αK145 (Figure 6-figure supplement 2).

Discussion

Added in the revised version: After diffusing to the target, the only agonist-dependent rearrangements (energy changes) in receptor activation are in catch and hold, A-C⇄ACLA⇄ACHA (Figure 1B, bottom) (Auerbach, 2024; Jadey & Auerbach, 2012). By comparing binding energies calculated in silico with those measured in vitro we were able to identify in the MD trajectories the hold end states, ACLA and ACHA, plus an intermediate configuration not detected in electrophysiology experiments. Three conspicuous hold rearrangements discussed below are: i) a cis→trans rotation of the agonist (flip), ii) an up→down displacement of loop C (flop), and iii) reductions in pocket volume, water and fluctuations (fix).

Removed in the revised version: In diffusion-catch-hold activation of AChRs (; Fig. 1B), the binding free energy increase in the last step generates the force that initiates the receptor isomerization. We identified in MD simulations (Fig. 3) the hold end states, ACLA (m1) and ACHA, (m3), as well as an intermediate configuration (m2). Our approach to understanding the conformational dynamics of hold is to compare calculated with experimental binding free energies, ΔGLA and ΔGHA. There are many experimental measurements for many different AChR agonists and several binding sites (Indurthi & Auerbach, 2023; Jadey & Auerbach, 2012; Purohit et al., 2014), but here we describe the comparison for only 4 agonists and only the wt α-δ site.

Earlier version: The alignment between experimental and calculated free energies was closer for ACh and CCh than for Ebt and Ebx (Table 1). The match improved when free energy ratios rather than absolute values were compared. Individual in silico values overestimated experimental ones by approximately the same factor (different for each ligand) that, however, cancelled in the ratio. We do not know the origin of this factor but speculate that it could be caused by errors in ligand parameterization. Regardless, the results show that agonist efficiency (η) can be estimated from only agonist and receptor structure (one PDB file), and that η is a useful metric for associating structure with function.

Revised version: There was reasonably good agreement between calculated and experimental binding free energies for ACh and CCh, but the match was excellent when free energy ratios (efficiencies) rather than absolute values were compared (Figure 4A). Absolute in silico values overestimated experimental ones by approximately the same factor (different for each ligand) that cancelled in the ratio. We do not know the origin of this factor but speculate that it is caused by errors in ligand parameterization. Regardless, the results show that agonist efficiency can be a useful metric for associating structure with function, and that efficiency can be estimated from structure alone.

Added in the revised version: Below we consider the structural elements associated with flip, flop and fix, namely the agonist (cis→trans), loop C (up→down) and the pocket (large→small, et al). MD simulations of the α–δ AChR neurotransmitter site suggest the following approximate and jittery generic sequence of rearrangements in ACLA→ACHA (Video 1). The stable states are bold.

Added in the revised version: ACLA (m1): agonist cis, loop C up, pocket wide/wet/wobbly,

Added in the revised version: …agonist leaves cis and starts to rotate about the N+-αW149 fulcrum,

Added in the revised version: …loop C starts to move down (towards the fulcrum).

Added in the revised version: Intermediate (m2): agonist half-rotated, tail stabilized by αY198,

Added in the revised version: …agonist rotates to trans (little change at the fulcrum), tail not secured in the pocket,

Added in the revised version: loop C is fully down.

Added in the revised version: ACHA (m3): agonist trans, tail secured by VdW contacts and H-bonds, pocket small/dry/rigid.

Added in the revised version: Agonist

Added in the revised version: With all agonists and in all simulations, the rotation of the agonist about the N+-αW149 fulcrum begins at the start of the ACLA→ACHA transition. Movements of bound ligands have been identified in other proteins. In AMPA-type glutamate receptors, the neurotransmitter glutamate binds in either of 2 poses, crystallographic or inverted (Yu et al., 2018). In one MD simulation these interconverted, suggesting the bound ligand is not rigidly fixed. Although the alternative poses were not associated with binding energies, the binding cleft was more closed in the crystallographic (presumably HA) versus inverted (possibly LA) orientation. The re-orientation of a bound glutamate in AMPA receptors is possibly related to the flip of the agonist in AChRs.

Added in the revised version: Multiple ligand positions in binding pockets have also been reported in Glycine receptors (Yu et al., 2014), 5-HT3 receptors (Basak et al., 2020), δ-opioid receptors (Shang et al., 2016), HIV-1 protease (Klei et al., 2007) and transthyretin (Klabunde et al., 2000). In enzymes, bound substrate movement is apparent in hexokinase (Bennett & Steitz, 1978) and alcohol dehydrogenase (Plapp, 2010). It is not unusual for movement of a bound ligand to trigger an otherwise unfavorable local rearrangement that drives protein function (an ‘induced fit’; (Richard, 2022). Not only can an untethered ligand come and go, it can also reorient itself dramatically to change binding energy and, hence, the local environment that couples to the full allosteric transition.

Added in the revised version: In AChRs, such ligand reorientation cannot pertain to all agonists. Some with high potency are perfectly symmetrical (decamethonium) and some lack a tail (tetramethylammonium). Although we have not yet identified the conformational dynamics that determine ΔGHA, factors other than the flip certainly contribute, and in proportions that are likely to be agonist dependent (see below).

Removed in the revised version: In the ACLA→ACHA simulations, two conspicuous binding site rearrangements were a cis→trans somersault of the agonist (a flip) and an up→down droop of loop C (a flop). The loop C movement is well known, but its functional role has been puzzling. The extent of the flop is correlated with agonist potency in ACh binding proteins (Basak et al., 2020; Du et al., 2015; Polovinkin et al., 2018), but voltage-clamp fluorometry of other receptors shows it to be similar with agonists and antagonists (Chang & Weiss, 2002; Munro et al., 2019; Pless & Lynch, 2009). In AChRs, truncation of loop C prevents activation by agonists but has no effect on unliganded activation, but agonist sensitivity is maintained if αY190 and αY198 are intact (Purohit & Auerbach, 2013). These results suggest that the downward displacement of loop C in AChRs is mainly associated with the generation of strong binding energy.

Removed in the revised version: Perhaps the most puzzling result regarding loop C is that the agonist association rate constant kon is slower to ACLA (up position) than to ACHA (down position) (Grosman & Auerbach, 2001; Nayak & Auerbach, 2017). Moreover, whereas kon to ACLA is correlated with agonist potency (the dissociation rate constant koff is approximately the same for all agonists), kon and koff to ACHA both are agonist dependent. The flop of loop C in hold removes a barrier to make kon approximately diffusion-limited, and also reduces discrimination between ligands (molecular recognition.

Removed in the revised version: Loop C droop resembles the ‘clamshell closure’ apparent at other receptors binding site (Chen et al., 2017; Hansen et al., 2005; Hibbs & Gouaux, 2011). However, in the light of the above experimental results, this label is misleading because ‘closing’ (or ‘capping’) a binding pocket is expected to slow rather than speed kon. For this reason, we use less functionally loaded words to describe the downward displacement of loop C, namely a ‘flop’ or ‘droop’. In ACLA→ACHA loop C indeed moves closer to αC of αW149, to cover the pocket and create higher affinity. However, it does not do so by imposing a steric barrier. We did not study the catch rearrangement and do not speculate about the nature of the barrier that limits the agonist association rate constant to CLA that has loop C in an up position.

Removed in the revised version: Importantly, loop C flop repositions αY190, the main determinant of both LA and HA binding energies (Bruhova & Auerbach, 2017; Purohit et al., 2014). In ACLA→ACHA, this side chain fills the slot adjacent to αY93 that otherwise allows the cis orientation in ACLA (Fig. S5). Further, flop moves αY190 closer to a salt bridge between αK145-αD200, and compacts the space between αW149 and N+ (Fig. 6). Overall, the simulated structures suggest that loop C flop increases binding energy both by repositioning αY190 near N+ and by allowing the cis-trans flip of the ligand. Rather than closing a door, the downward movement of loop C rearranges the furniture.

Removed in the revised version: We interpret cis-trans rotation of the agonist about the N+-αW149 fulcrum (Fig. 5) with caution. Although the trans orientation is apparent in desensitized AChR cryo-EM structures, the cis orientation exists only in silico. Notwithstanding this limitation, in ACLA→ACHA the agonist’s flip into the binding pocket and the repositioning αY190 by loop C flop were observed consistently in all simulations with all agonists, and rationalize the increase binding energy in the hold transition. We speculate that the agonist flip triggers the rest of the hold restructuring but cannot rule out flop promoting flip or synchronous rearrangements. We have been unable to establish clearly the flip-flop temporal sequence.

Removed in the revised version: For the agonists we examined, the somersault positions the tail deep into the binding cavity where it interacts via a structural water with the complementary subunit. It appears that agonists rotation increases binding free energy by displacing water and increasing vdW interactions and protein stability. A hydrophobic binding pocket is associated with functional activation in other ion channels (Furukawa et al., 2005; Sigel & Steinmann, 2012) and is observed in many proteins including β2-adrenergic receptor, opioid receptors, and HIV-1 protease (Huang et al., 2015; Rasmussen et al., 2011; Wlodawer & Erickson, 1993). However, these effects on the AChR pocket cannot be the only basis for the higher affinity ACHA because several potent agonists lack a tail, for example tetramethylammonium. We hypothesize that both the reorganization of the aromatic groups in the pocket around N+ and binding site compaction/de-wetting contribute to stronger binding, in proportions that remain to be determined. Regardless, the agonist’s somersault shows that mobility is one feature that distinguishes a ligand from a covalently linked side chain (Indurthi & Auerbach, 2023). Not only can an unconnected agonist come and go, it can also reposition itself dramatically to induce a conformational change of binding pocket to affect a change in binding energy.

Removed in the revised version: Movements analogous to flip and flop have been identified in other proteins. Multiple ligand positions in a binding pocket have been documented in 5-HT3 receptors (Basak et al., 2020), δ-opioid receptors (Shang et al., 2016), HIV-1 protease (Klei et al., 2007) and transthyretin (Klabunde et al., 2000). In enzymes, substrate motion has been observed in hexokinase (Bennett & Steitz, 1978) and alcohol dehydrogenase (Plapp, 2010). Regarding flop, substrate-induced loop displacements are known to reposition residues to enhance catalysis, for example in chymotrypsin (Ma et al., 2005) and triosephosphate isomerase (Brown & Kollman, 1987; Nickbarg et al., 1988). The agonist flip and loop C flop in AChRs are additional examples of ligand mobility and loop motions undergirding function.

Earlier version: The PCA plots (Fig. 3) show 3 structural populations, one of which (m2) occurs in time between m1 (ACLA) and m3 (ACHA). In m2, the agonist tail interacts with αY198, possibly to guide the clockwise rotation of the ligand and reduce the probability of a return to the cis orientation. ACh was the only agonist to rotate counterclockwise, perhaps because it lacks the appropriate functional group in its tail. The MM-PBSA calculations indicate that ΔGm2 (that does not have an experimental correlate) is intermediate between ΔGLA and ΔGHA for ACh and CCh, but least stable for Ebt and Ebx (Table 1). We do not have an explanation for this difference. However, we observe m1→m2, m2→m1, m2→m3 and m3→m1 transitions in the trajectories (Video 1), so the depth of the m2 energy well (relative to the barriers that separate it from m1 and m2) contributes to agonist efficacy and efficiency.

Revised version: The PCA plots (Figure 3) show 3 structural populations, one of which (m2) occurs in time between m1 (ACLA) and m3 (ACHA) (Figure 3-figure supplement 2). In m2, the agonist’s tail (CCh, Ebt and Ebx) interacts with αY198, possibly to guide the clockwise rotation of the ligand and reduce the probability of a return to the cis orientation. The PBSA calculations indicate that ΔGm2 (that does not have an experimental correlate) is intermediate between ΔGLA and ΔGHA for ACh and CCh, but least stable for Ebt and Ebx (Figure 4-Source Data 1). We do not have an explanation for this difference.

Added in the revised version: Loop C

Added in the revised version: A number of observations place focus on the role of the downward displacement of loop C in receptor activation. In ACh binding proteins, its extent correlates with agonist potency (Basak et al., 2020; Du et al., 2015; Polovinkin et al., 2018) even if voltage-clamp fluorometry of other receptors shows it to be similar with agonists and antagonists (Chang & Weiss, 2002; Munro et al., 2019; Pless & Lynch, 2009). Simulations of the apo AChR indicate that in the absence of agonists loop C has flopped partially, narrowing the aqueous connection between the bath and the pocket (Figure 2-figure supplement 2C).

Added in the revised version: Given the flop in loop C, it is puzzling that electrophysiology measurements show the agonist association rate constant (kon) is substantially faster to CHA (narrow entry, loop C down) than to CLA (wide entry, loop C up), as well as more correlated with agonist potency (Grosman & Auerbach, 2001; Nayak & Auerbach, 2017). Docking results offer a clue to the apparent contradiction that narrowing an entrance speeds association up to the diffusional limit (4 x109 M−1s−1). All 4 agonists dock in the cis orientation to CLA (wide) versus in the trans orientation to CHA (narrow) (Figure 2B). Hence, despite the wider aperture, trans is disfavored and lower affinity cis is preferred. We hypothesize that binding in the trans orientation is prohibited in the apo pocket, and that loop C flop (and other rearrangements) both narrow the entrance and remove this prohibition. Accordingly, kon with loop C up (LA) is slower than with loop C down (HA). We do not know the structural element(s) that prohibits trans with a wide pocket. Any or all of the constellation of rearrangements in unliganded CLA→CHA could contribute.

Added in the revised version: As an aside, the inward displacement of loop C toward the pocket center resembles the ‘clamshell closure’ apparent at other receptor binding sites (Chen et al., 2017; Hansen et al., 2005; Hibbs & Gouaux, 2011). However, in light of the kon experimental results, using ‘closure’ (or ‘capping’) to describe this motion is misleading because it implies putting a lid on the pocket that would reduce, not increase, the agonist association rate constant. For this reason, we use functionally neutral words ‘flop’ and ‘up/down’ for loop C displacement and position.

Added in the revised version: AChRs can turn on without a loop C. Truncation of loop C eliminates activation by ACh but has almost no effect on unliganded activation (Purohit & Auerbach, 2013). Interestingly, leaving only loop C residues αY190 and αY198 in place saves activation by ACh. Hence, repositioning of one (or both) of these in up→down is sufficient to maintain activation by the neurotransmitter. Other agonists were not investigated.

Added in the revised version: αY190 is our interest. An Ala substitution here results in a large loss of ACh binding energy with about half coming from the removal of the hydroxyl group (Purohit et al., 2014). In contrast, the mutations αY198F and αY93F have essentially no effect on ACh binding energy. These results suggest that repositioning αY190 consequent to loop C flop is important for activation by ACh (Figure 6). The αY190 ring fills the slot that opens next to αY93 when the agonist exits cis, preventing a return (Figure 5-figure supplement 1). However, the distance between the αY190 ring and N+ does not change significantly so it is unlikely that the deployment of αY190 increases binding energy via an interaction with the ring (Figure 6-table supplement 1).

Added in the revised version: Perturbation of the salt bridge residue αD200 by the αY190 hydroxyl has been proposed to trigger the AChR allosteric transition (Beene et al., 2002; Gay & Yakel, 2007; Mukhtasimova et al., 2005; Padgett et al., 2007; Pless et al., 2011). In support of this idea, the mutation αK145A reduces ACh (and choline) binding energy substantially but has no effect when the αY190 hydroxyl is absent (Bruhova & Auerbach, 2017). In the simulations, the downward displacement of loop C positions the αY190 hydroxyl to within ~5.0 Å of αD200 and αK145 (Figure 6-table supplement 1). We sought, but did not find, structural water in this gap. Although there was no correlation between these separations and the degree of loop C displacement, the gap was smaller with high- (ACh and CCh) versus low-efficiency agonists (Ebt and Ebx). In summary, some evidence supports the above proposal, but so far, the simulations have not illuminated a mechanism. Substrate-induced loop displacements are known to deploy residues to enhance catalysis, for example in chymotrypsin (Ma et al., 2005) and triosephosphate isomerase (Brown & Kollman, 1987; Nickbarg et al., 1988). Repositioning αY190 in the AChR ACLA→ACHA transition could be analogous.

Added in the revised version: Although attention on loop C has focused on the role of the flop, it is also important to consider the functional consequences of the up position. The structural elements of the up (wide cleft) configuration that prohibit agonist docking into the trans orientation have not been identified. A key requirement for allostery is that the ligand bind more strongly to one of the alternative conformations of the orthosteric site, so attention must also be given to the conformational dynamics of the ‘catch rearrangement, A-C⇄ACLA. In both GABAA (Jones et al., 2001; Mortensen et al., 2010) and ACh receptors (Jadey & Auerbach, 2012) the LA agonist association rate constant is correlated with potency rather than diffusion constant, indicating that the formation of ACLA requires an energy (structure) change.

Added in the revised version: Pocket

Added in the revised version: The αW149-N+ hinge remains relatively unchanged in the hold transition (Figure 6, Figure 6-table supplement 1). Hence, we hypothesize the bulk of the extra favorable agonist binding energy in ACHA is not generated here but somewhere else, cis→trans. For the agonists we examined, the rotation places the tail into the main binding cavity (Figure 6-figure supplement 1) (Video 1). The trans orientation is secured by VdW interactions (possibly caused by further downward loop C displacement and dehydration) and a hydrogen bond with a structural water bonded to the loop E backbone (in the complementary subunit). The net result to the overall pocket is a reduction in volume (Figure 4C) and increase in stability of loops C and F (Figure 3).

Added in the revised version: In the simulations, these two structural changes, distributed over the large surface area of the cavity, were robust and consistent for all 4 agonists. It is possible that the increase in binding free energy in hold derives mainly from restructuring of the pocket rather than from a change in the N+-aromatic interactions or, indeed, any one local ligand-protein interaction. Compact, hydrophobic binding pockets are associated with functional activation in other ion channels (Furukawa et al., 2005; Sigel & Steinmann, 2012) and observed in other proteins including β2-adrenergic receptors, opioid receptors, and HIV-1 protease (Huang et al., 2015; Rasmussen et al., 2011; Wlodawer & Erickson, 1993).

Added in the revised version: After hold, the next step In AChR gating is the rearrangement of the extracellular domain (Auerbach, 2024; Purohit et al., 2013). In related receptors, this domain becomes more compact and stable (a rearrangement called ‘unblooming’) (Sauguet et al., 2014; Taly et al., 2009), structural changes that echo those in hold at the AChR neurotransmitter site. It is therefore possible that the ACLA→ACHA rearrangement of the neurotransmitter site nucleates the subsequent restructuring of the extracellular domain. Extending this speculation, the structural correlates of energy change in gating might at first spread in contact area (from N+, to pocket, to extracellular domain) and then refocus (from transmembrane domain, to gate region, to pore water).

Removed in the revised version: MD simulations did not reveal an interaction between the αK145-αD200 salt bridge distance and the degree of loop C displacement (Cheng et al., 2006). However, this distance discriminate between η classes by decreasing ACLA→ACHA in high (ACh and CCh) but not in low (Ebt and Ebx) η classes, consistent with experimental results (Indurthi & Auerbach, 2023). All 4 agonists are high efficacy, so this result suggests that perturbation of this salt bridge is not a necessary requirement for the change from LA to HA. After loop C flop, the repositioned αY190 connects N+ to the salt bridge, suggesting that this interaction (rather than bridge tightening) could help trigger the gating transition (Beene et al., 2002; Gay & Yakel, 2007; Mukhtasimova et al., 2005; Padgett et al., 2007; Pless et al., 2011). The simulation did not reveal the conformational changes (forces) that drive the next step of the isomerization (ECD restructuring). However, the observation that for all agonists, interactions with the complementary subunit and pocket dehydration terminated (at 200 ns) the hold rearrangement lead us to speculate these may contribute.

Earlier version: Loop E (connecting β5 and β6 strands) has been implicated in ligand discrimination (Basak et al., 2020; Muroi et al., 2009; Pless & Lynch, 2009). The double mutation F106L+S108C in the β2 subunit of the α4β2 nicotinic AChR subtype significantly decreases the action of Ebt without altering that of ACh (Tarvin et al., 2017). Loop E residues are responsible for the binding energy difference of α-γ versus α-Δ interfaces because hydrogen bonds influence the position and participation of γW55 (Nayak et al., 2016), and for reducing instability of the α-Δ site consequent to loop C mutations (Purohit & Auerbach, 2013). MD simulations with more agonists and more binding sites are needed to elucidate the role of loop E in the ACLA→ACHA transition.

Revised version: In the complementary δ subunit (Figure 1D), loop E (β5-β6 strands) has been implicated in ligand discrimination (Basak et al., 2020; Muroi et al., 2009; Pless & Lynch, 2009). The double mutation F106L+S108C in the β2 subunit of the α4β2 nicotinic AChR subtype significantly decreases the action of Ebt without altering that of ACh (Tarvin et al., 2017). In addition, loop E residues are responsible for the binding energy differences α−γ versus α−δ (Nayak et al., 2016) and reduce site instability caused by loop C mutations (Vij et al., 2015). Although our simulations do not shed light on these matters, the temporal sequence flip-flop-fix suggests that perturbation of the loop E backbone could also play a role in downstream activation of the receptor.

Earlier version: Agonists belong to efficiency (η) classes in which ΔGLA/ΔGHA is the same for all members. So far 5 η-classes have been identified experimentally in adult-type AChRs, indicating that are 5 pairs of ACLA/ACHA structures (Indurthi & Auerbach, 2023). Although in the simulations loop F movements and salt bridge distances were agonist dependent, with only 4 ligands we cannot associate these rearrangements with η classes.

Revised version: Agonists belong to efficiency classes in which ΔGLA/ΔGHA is the same for all members. So far 5 such classes have been identified experimentally in adult-type AChRs, indicating that also are 5 pairs of ACLA/ACHA structural classes (Indurthi & Auerbach, 2023). Although in the simulations loops C and F movements, salt bridge distances were agonist dependent, with only 4 ligands we cannot associate any these rearrangements with efficiency classes. Additional simulations using more agonists and with binding site mutations may lead to better understanding of the dynamics that underpin molecular recognition in catch and protein animation in hold, and perhaps reveal the structural basis of their linkage (Auerbach, 2024).

Added in the revised version: Implications

Added in the revised version: Several aspects of this work may have general implications. The LA and HA end states of hold emerged from simulations of a single protein structure and were identified by comparing experimental with calculated binding energies. If general, this approach could be useful because cryo-EM structures of ligands bound with LA are lacking. Agonist efficiency was calculated accurately from structure alone. This, too, could be useful because efficiency allows affinity and efficacy to be calculated from each other (Indurthi & Auerbach, 2021). In AChRs the up→down movement of loop C flop does not impose a lid on the pocket. Rather, the up position serves to enforce catch (ligand recognition) by blocking directly binding in the trans orientation, and the down position serves to stabilize the ACHA pocket and, perhaps, to deploy αY190. It is possible that the ‘clamshell closures’ apparent in other receptors also serve to activate for reasons other than imposing a steric barrier. Finally, it appears that the movement of the agonist, bound with low affinity, is the initial event in gating. If confirmed, we think it both fitting and beautiful that the motion of a structural element added to the system from the outside jumpstarts the allosteric transition.

Removed in the revised version: The key requirement for allostery is that ligand binding energy be weaker to the resting versus active conformation of the target. Natural selection must generate both low and high affinity binding sites, so our attention must be on restructuring in both AC=ACLA (catch) and ACLA=ACHA (hold). What prevents loop C flop in apo? What barrier(s) slows kon to apo? Where is the agonist in the encounter complex? What forces promote the cis orientation? What is the role of water in agonist orientation, both cis and trans? What are the mechanisms and consequences of a wider salt bridge, apo versus ACLA? What structural classes are associated with efficiency classes? The correlation in the free energy changes in catch and hold that entangles binding with gating implies a corresponding correlation in structural change. We anticipate that additional MD simulations of catch, hold, and post-hold domain rearrangements, for more agonists and with binding site mutations, will lead to a more complete understanding of how agonists turn on the AChR.

Materials and Methods

Hardware and Software

Earlier version: All computational studies were carried out on a Linux Ubuntu 20.04 based workstation with 126 GB RAM, 64 CPU cores, and two RTX A5000 24 GB GPU cards, as well as an Ubuntu 20.04 based server equipped with 132 GB RAM, 96 CPU cores, and two RTX A5000 24 GB GPU cards. The software utilized includes the Maestro release 2022-2 graphical user interface (GUI) of the Schrödinger software suite (BioLuminate, Schrödinger, LLC, NY, USA), Amber22 (Case et al., 2005), VMD 1.9.3 (Humphrey et al., 1996), ChemDraw Professional 15.1 (Purushotham et al., 2022), ChimeraX (Pettersen et al., 2021), and PyMOL (PyMOL Molecular Graphics System, Version 2.0, Schrödinger, LLC).

Revised version: All computational studies were carried out on a Linux Ubuntu 20.04 based workstation with 126 GB RAM, 64 CPU cores, and two RTX A5000 24 GB GPU cards, as well as an Ubuntu 20.04 based server equipped with 132 GB RAM, 96 CPU cores, and two RTX A5000 24 GB GPU cards. The software utilized includes the Maestro release 2022-2 graphical user interface (GUI) of the Schrödinger software suite (BioLuminate, Schrödinger, LLC, NY, USA), Amber22 (Case et al., 2005), VMD 1.9.3 (Humphrey et al., 1996), ChemDraw Professional 15.1 (Purushotham et al., 2022), ChimeraX (Pettersen et al., 2021), and PyMOL (PyMOL Molecular Graphics System, Version 2.0, Schrödinger, LLC). Protein preparation. Because agonist energy changes are approximately independent of energy changes in the protein outside the binding pocket, for the simulations we removed the transmembrane domain and simulated only α−δ subunit extracellular domain dimers. The starting structure for the hold transition was resting C (PDB: 6UWZ) after toxin removal and is essentially the same as the apo structure (Zarkadas et al., 2022) (Figure 1C). The 4 agonists we examined were acetylcholine (ACh), carbamylcholine (CCh), epibatidine (Ebt) and epiboxidine (Ebx) (Figure 2A). These are of similar size (MW 187, 183, 209 and 178, respectively) but have different efficiencies (0.50, 0.52, 0.42, and 0.46, respectively) (Indurthi & Auerbach, 2023).

Removed in the revised version: Protein preparation

Protein preparation was conducted using the Protein Preparation Wizard (PPW) (Purushotham et al., 2022) and OPLS4 force field. Bond orders were assigned to atoms based on the CCD database, with zero bond orders given to bonded metals. Missing hydrogen atoms were inserted, and disulfide bonds were placed appropriately. During preprocessing, structural refinement was performed with Maestro’s Prime utility, optimizing and adding any missing side chains and residues. Missing residues were formulated by combining SEQRS information from the PDB files and optimized using Prime’s ePLOP module. The ionization state of all heterogeneous groups, including ligands, bound metals, and charged amino acids, were determined using the EpiK tool at a pH range of 7.0 ± 2.0. Subsequent force field minimization allowed for water sampling and basic restraint minimization, during which only hydrogen atoms could be freely reduced.

Earlier version: The hydrogen bonding (H-bond) network underwent significant modification. The terminal chi angle for Asn, Gln, and His was sampled through 180-degree flips, altering their spatial H-bonding capabilities without affecting the fit to electron density. The neutral and protonated states of His, Asp, and Glu, as well as two His tautomers, were sampled. Hydrogens on hydroxyls and thiols were also analyzed to enhance the H-bond network. Finally, the ProtAssign method was employed in two modes: a "regular" mode, which completed in seconds, and an "exhaustive" mode that considered many more states and could take minutes or hours, depending on the H-bond network complexity.

Revised version: The hydrogen bonding (H-bond) network underwent significant modification. The terminal chi angle for residues Asn, Gln, and His was sampled through 180-degree flips, altering their spatial H-bonding capabilities without affecting the fit to electron density. The neutral and protonated states of His, Asp, and Glu, as well as two His tautomers, were sampled. Hydrogens on hydroxyls and thiols were also analyzed to enhance the H-bond network. Finally, the ProtAssign method was employed in two modes: a “regular” mode, which completed in seconds, and an “exhaustive” mode that considered many more states and could take minutes or hours, depending on the H-bond network complexity. Ligand Preparation. The two-dimensional SDF structures of carbamylcholine (CCh), acetylcholine (ACh), epibatidine (Ebt) and epiboxidine (Ebx) were retrieved from the PUBCHEM database. These agonists were chosen because they have a similar size and represent different efficiency classes (Indurthi & Auerbach, 2023). LigPrep (LigPrep, Maestro 11.2.014, Schrödinger LLC) was utilized to construct ligands, produce stereoisomers, and generate tautomers using the OPLS46 force field. Additionally, the Epik module was employed to desalt and protonate the ligands at pH 7.0. Default values were applied to the remaining parameters, and LigPrep was also used to prepare the ligands ACh, CCh, Ebt, and Ebx, with the option “retain specified chiralities” engaged.

Removed in the revised version: Ligand Preparation

Removed in the revised version: The two-dimensional SDF structures of carbamylcholine (CCh), acetylcholine (ACh), epibatidine (Ebt) and epiboxidine (Ebx) were retrieved from the PUBCHEM database. LigPrep (LigPrep, Maestro 11.2.014, Schrödinger LLC) was utilized to construct ligands, produce stereoisomers, and generate tautomers using the OPLS46 force field. Additionally, the Epik module was employed to desalt and protonate the ligands at pH 7.0. Default values were applied to the remaining parameters, and LigPrep was also used to prepare the ligands ACh, CCh, Ebt, and Ebx, with the option "retain specified chiralities" engaged.

In terms of ligand state assessment, Epik can calculate a penalty to quantify the energy cost of forming each state in solution based on the Hammett and Taft methods. This EpiK state penalty is entirely compatible with the GlideScore used for docking, as it is computed in kcal/mol units. This compatibility facilitates the examination of how the EpiK state penalty affects the GlideScore. The DockingScore in Glide represents the total of the GlideScore and the Epik state penalty and served as the basis for final ranking and enrichment evaluation. Furthermore, Epik offers a technique for handling metal binding states, involving a change in the pH range within which the state is generated. This approach underscores the comprehensive considerations made in preparing the ligands for docking and subsequent analyses.

Grid Generation and Molecular Docking

The receptor grid complex of 6UWZ.pdb was generated after removing alpha-bungarotoxin from the binding pocket using Glide (Glide, Maestro 11.2.014, Schrödinger LLC). A virtual grid box with dimensions 20Å × 20Å × 20 Å was formed, centered within the C and F loops in the α–δ binding site of the energy-minimized resting state. This grid box was used to generate a receptor grid, and the default values were assigned for other parameters. Docking was performed using the ligand docking wizard built into the Glide module of Schrödinger, with the default scaling factor and partial charge cutoff set to 0.80 and 0.15, respectively. The docking procedure consisted of two stages, extra precision (XP) followed by post-processing with prime MM-GBSA. The best dock poses were selected based on dock score and MM-GBSA energy for further investigation.

Molecular Dynamics Simulations

Earlier version: MD simulations were conducted using the Amber22 software suite within the Amber module as detailed elsewhere (Singh et al., 2021; Srivastava et al., 2022). Ligand topology was prepared via the Leap module in AmberTools22, employing force field ff19SB (Tian et al., 2020) and the Generalized Amber Force Field (GAFF) (Vassetti et al., 2019). The AM1-BCC strategy and GAFF2 were used to assign atomic charges and additional parameters. Before simulation, the systems were minimized and equilibrated following a five-step minimization process, with each step involving 10,000 energy minimization steps. Systems were heated under NVT ensembles in two consecutive runs from 0–300 K for 1 ns each, followed by a 1 ns simulation at 300 K and 1 bar pressure under the NPT ensemble for equilibration. An additional 5 ns equilibration was performed before the 200 ns production MD for each system, conducted under periodic boundary conditions (PBCs) and utilizing an NPT ensemble with 300 K and 1 bar pressure. Long-range electrostatic interactions were treated using the particle mesh Ewald method (PME), including a direct space cutoff of 12.0 Å and van der Waals interactions. In production MD, a time step of 2.0 fs was employed, and the SHAKE algorithm was used to keep bond lengths at equilibrium. Isobaric (NPT) conditions were maintained using the Berendsen barostat, with temperature monitored by a Langevin thermostat. Coordinates from the production MD were recorded in the trajectory file every 20 ps, resulting in a minimum of 10,000 frames. Trajectories were analyzed using amber-tools CPPTRAJ (Roe & Cheatham, 2013), and further analysis and figure generation were completed with VMD, Pymol, Xmgrace and ChimeraX. To ensure reproducibility, MD trajectories were generated in triplicate.

Revised version: MD simulations were conducted using the Amber22 software suite within the Amber module as detailed elsewhere (Singh et al., 2021; Srivastava et al., 2022). Ligand topology was prepared via the Leap module in AmberTools22, employing force field ff19SB (Tian et al., 2020) and the Generalized Amber Force Field (GAFF) (Vassetti et al., 2019). The AM1BCC strategy and GAFF2 were used to assign atomic charges and additional parameters. Before simulation, the systems were minimized and equilibrated following a five-step minimization process, with each step involving 10,000 energy minimization steps. Systems were heated under NVT ensembles in two consecutive runs from 0–300 K for 1 ns each, followed by a 1 ns simulation at 300 K and 1 bar pressure under the NPT ensemble for equilibration. An additional 5 ns equilibration was performed before the 200 ns production MD for each system, conducted under periodic boundary conditions (PBCs) and utilizing an NPT ensemble with 300 K and 1 bar pressure. Long-range electrostatic interactions were treated using the particle mesh Ewald method (PME), including a direct space cutoff of 12.0 Å and van der Waals interactions. In production MD, a time step of 2.0 fs was employed, and the SHAKE algorithm was used to keep bond lengths at equilibrium. Isobaric (NPT) conditions were maintained using the Berendsen barostat, with temperature monitored by a Langevin thermostat. Coordinates from the production MD were recorded in the trajectory file every 20 ps, resulting in a minimum of 10,000 frames. Trajectories were analyzed using amber-tools CPPTRAJ (Roe & Cheatham, 2013), and further analysis and figure generation were completed with VMD, Pymol, Xmgrace and ChimeraX. To ensure reproducibility, MD trajectories were generated in triplicate.

Cluster Analysis

Cluster analysis of the ligand was performed using TCL scripts within the Visual Molecular Dynamics (VMD) software, as described elsewhere (Mittal, Tonk, et al., 2021). In brief, this script automates the identification of frames from the most stable conformations based on PCA minima. Using these frames as input, the algorithm employed RMSD values of ligand positions to cluster similar orientations. Specifically, clusters with an RMSD of ≤ 1Å were considered to share a similar ligand orientation. Frames that did not fit into these top clusters were automatically excluded by the script.

Principal Component Analysis (PCA) and Free Energy Landscape (FEL)

Earlier version: PCA identifies collective motions in atomic-level simulations of macromolecules (Mittal, Srivastava, et al., 2021; Srivastava et al., 2022). It emphasizes significant patterns by reducing noise in MD simulations, revealing essential dynamics underlying conformational changes. All equilibrated MD simulation trajectories were analyzed by PCA computations using the CPPTRAJ algorithm in Amber tools.

Revised version: PCA identifies collective motions in atomic-level simulations of macromolecules (Mittal, Srivastava, et al., 2021; Srivastava et al., 2022). It emphasizes significant patterns by reducing noise in MD simulations, revealing essential dynamics underlying conformational changes. All equilibrated MD simulation trajectories were analyzed by PCA computations using the CPPTRAJ algorithm using Amber tools.

Removed in the revised version: Prior to PCA, least-square fitting of the protein’s Cα atoms was carried out with respect to a reference structure (average protein coordinates) to eliminate rotational and translational degrees of freedom in the simulation box. A positional covariance matrix C (of dimensions 3N × 3N) was constructed using the Cα atom coordinates, and this matrix was diagonalized to produce eigenvectors (directions of motion) and eigenvalues (magnitudes). The elements of the covariance matrix C were obtained using:

Earlier version: where Xi and Xj are the Cartesian coordinates of the Cα atoms, and N is the number of Cα atoms. The ⟨⟩ notation represents the ensemble average of atomic positions in Cartesian space. The resulting principal components (eigenvectors) were sorted by the total motion they captured, generating 3N eigenvectors. The free energy landscapes (FELs; FIg.3) were created using the "g_sham" module in GROMACS, based on the probability distribution of the top two principal components (PCs). The FEL plot aids in visualizing possible conformations and corresponding energy levels. The free energy values (G) were calculated as

Revised version: Prior to PCA, least-square fitting of the protein’s Cα atoms was carried out with respect to a reference structure (average protein coordinates) to eliminate rotational and translational degrees of freedom in the simulation box. A positional covariance matrix C (of dimensions 3N × 3N) was constructed using the Cα atom coordinates, and this matrix was diagonalized to produce eigenvectors (directions of motion) and eigenvalues (magnitudes). The elements of the covariance matrix C were obtained using: where Xi and Xj are the Cartesian coordinates of the Cα atoms, and N is the number of Cα atoms. The ⟨⟩ notation represents the ensemble average of atomic positions in Cartesian space. The resulting principal components (eigenvectors) were sorted by the total motion they captured, generating 3N eigenvectors. The free energy landscapes (FELs; Figure 3) were created using the “g_sham” module in GROMACS, based on the probability distribution of the top two principal components (PCs). The FEL plot aids in visualizing possible conformations and corresponding energy levels. The free energy values (G) were calculated as where k is the Boltzmann constant, T is the absolute temperature of the simulated system (300 K), Ni is the population of bin i, and Nmax is the population of the most-populated bin.

Removed in the revised version: where k is the Boltzmann constant, T is the absolute temperature of the simulated system (300 K), Ni is the population of bin i, and Nmax is the population of the most-populated bin.

Binding Free Energy Calculation

Added in the revised version: Experimental binding free energies were estimated by using electrophysiology, as described elsewhere (Indurthi & Auerbach, 2023). Briefly, the constants KHA and L2 were estimated by using electrophysiology (dose-response curves, for example), and KHA was calculated using the thermodynamic cycle with L0 known a priori. The logs of KLA and KHA are proportional to ΔGLA and ΔGHA. The binding free energy of the protein-ligand docked complex (ΔGbind) was calculated in two stages, after docking (using MM-GBSA in the Prime module of Schrödinger) (Mittal, Tonk, et al., 2021; Purushotham et al., 2022) and after MD simulation on equilibrated MD trajectories (using MM-PBSA/-GBSA in AMBER22). MM-PBSA is particularly useful to compare binding affinities between different states and among agonists (Kollman et al., 2000; Rastelli et al., 2010; Sun et al., 2014). In our use, MM-PBSA computations included both enthalpic (ΔH) contributions from molecular interactions (van der Waals, electrostatics, solvation) and entropic (TΔS) contributions calculated using normal mode analysis. By capturing the dynamic interplay between the ligand and receptor as well as the influence of the surrounding solvent environment this approach enables computation of a Gibbs free energy (ΔG).

Removed in the revised version: The binding free energy of the protein-ligand docked complex (ΔGbind) was calculated in two stages, after docking (using MM-GBSA in the Prime module of Schrödinger) (Mittal, Tonk, et al., 2021; Purushotham et al., 2022) and after MD simulation on equilibrated MD trajectories (using MM-PBSA/-GBSA in AMBER22).

Earlier version: After docking, for the static docked pose of the protein-ligand complex the MM-GBSA binding energy calculation was performed using the OPLS4 force field and the variable-dielectric generalized Born (VSGB) model in the Prime MM-GBSA module in Schrödinger 2022. We used ΔGbind_dock to rank protein-ligand docked complexes, selecting those with the highest negative values and dock scores as the best ligands for further studies. After MD Simulation, in the MD-equilibrated system binding energies were calculated through the MM-PBSA/-GBSA protocol, implemented in AMBER22. This was performed over the most stable clustered poses extracted from the equilibrated trajectory at minima m1, m2, and m3 on the free energy landscape (FEL) over a 200 ns time frame. AMBER utilizes the conventional g_mmpbsa module for MM-PBSA calculations using MM-PBSA.py script. These methods calculate the energies of electrostatic interactions, van der Waals interactions, polar solvation, and non-polar solvation in the equilibrated trajectory using a single trajectory protocol (STP). The binding energy of each protein-ligand complex (ΔGbind) was calculated as

Revised version: After docking, for the static docked pose of the protein-ligand complex the MM-GBSA binding energy calculation was performed using the OPLS4 force field and the variable-dielectric generalized Born (VSGB) model in the Prime MM-GBSA module in Schrödinger 2022. We used ΔGbind_dock to rank protein-ligand docked complexes, selecting those with the highest negative values and dock scores as the best ligands for further studies. After MD Simulation, in the MD equilibrated system binding energies were calculated through the MM-PBSA/-GBSA protocol, implemented in AMBER22. This was performed over the most stable clustered poses extracted from the equilibrated trajectory at minima m1, m2, and m3 on the free energy landscape (FEL) over a 200 ns time frame. AMBER utilizes the conventional g_mmpbsa module for MM-PBSA calculations using MM-PBSA.py script. These methods calculate the energies of electrostatic interactions, van der Waals interactions, polar solvation, and non-polar solvation in the equilibrated trajectory using a single trajectory protocol (STP). The binding energy of each protein-ligand complex was calculated as ΔGbind=ΔGcomplex−ΔGprotein−ΔGligand where ΔGcomplex, ΔGprotein, and ΔGligand represent the absolute free energies of the complex, protein, and ligand systems, respectively. Additional terms, such as ΔH (enthalpy), T and ΔS (temperature and total solute entropy), ΔGgas (total gas phase energy), and ΔGsolv (solvation energy), contribute to the calculation.

Removed in the revised version: where ΔGcomplex, ΔGprotein, and ΔGligand represent the absolute free energies of the complex, protein, and ligand systems, respectively. Additional terms, such as ΔH (enthalpy), T and ΔS (temperature and total solute entropy), ΔGgas (total gas phase energy), and ΔGsolv (solvation energy), contribute to the calculation.

The Poisson-Boltzmann calculations were performed using an internal PBSA solver in Sander, while Generalized Born ESURF was calculated using ‘LCPO’ surface areas. The vibrational mode entropy contributions (TΔS) for protein-ligand interactions were calculated and averaged using normal-mode analysis in Amber. The entropy term provides insights into the disorder within the system, as well as how this disorder changes during the binding process. This value was added to the enthalpy term computed by the MM-PBSA method to determine the absolute binding free energy (ΔGbind). These comprehensive calculations allowed for a precise evaluation of the interactions between the protein and ligand, providing valuable insights into the binding process and its energetic characteristics.

Pocket Volume Calculation

To investigate the variations in binding site volume across all four agonist-bound cases, pocket volume analysis was carried out using POVME3.0 (Wagner et al., 2017). From each of the four systems, frames were extracted at consistent intervals with a stride of 2.0, culminating in representative sets of 5000 protein structures for every system. Before commencing the volume calculation, the trajectory was aligned, and frames were selected from VMD, serving as the initial input for this method. Subsequently, an inclusion region was defined, encapsulating all binding pocket conformations within the trajectory. This region’s definition involved building a sphere upon which the inclusion grid-points were computed, using different atoms of the ligand for this calculation.

Supporting information